## CSE 255 Problem Set 1: Collinear points

For this problem set, we'll be using the Jupyter notebook:

![](jupyter.png)

### In this problem set, you will write python3 code to find sets of collinear points given __ 2D points.

Definition of collinearity[1]: In geometry, collinearity of a set of points is the property of their lying on a single line. A set of points with this property is said to be collinear.

![](non-collinear-points.jpg)[2]

Here, points P,Q,R and A,R,B are collinear. However, points A,B,C are non-collinear.

For instance, if given points: {(1,1),(0,1),(2,2),(3,3),(0,5),(3,4),(5,6),(0,-3),(-2,-2)}

Sets of collinear points are: { {(0,1),(0,5),(0,-3)}, {(1,1),(2,2),(3,3),(-2,-2)} }. Note that the ordering of points within each set is irrelevant. Also, the order of sets does not matter. Besides, every set has to have atleast three points (single point or a pair of points always lie on a line). s

### TODO: Make changes as required
Outline of the problem set:

Part A: Load the data.

Part B: Remove duplicates.

Part C: Sort the data?s

Part D: Write logic to find collinear points.

Part E: Submit this notebook.

1. https://en.wikipedia.org/wiki/Collinearity
2. http://www.mathcaptain.com/geometry/collinear-points.html

## Points of concerns for TAs (keep adding stuff as you go along)

1. What should be the input format? Should we give them a file containing a bunch of 2D points?
2. Do we let them import any library they want or enforce it?


```python
## generate data points that go in the data input file for students

from itertools import product
import numpy as np

l = list(product(range(-400,400,np.random.randint(21)), range(-400,400,np.random.randint(20))))
```

```python
from pyspark import SparkContext, SparkConf
conf = SparkConf().setAppName("Collinear Points").setMaster("local[4]")
sc = SparkContext(conf=conf)
```

---
## Part A: Load the data (2 points)


Write a function that returns a list of numbers, such that $x_i=i^2$, for $1\leq i \leq n$. Make sure it handles the case where $n<1$ by raising a `ValueError`.

<!-- grade_id: squares, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()

def get_collinear(filename):
    """Compute the set of collinear points where each set of points is sorted"""
    rdd = sc.textFile(filename)
    # YOUR CODE HERE
    raise NotImplementedError()
```

Your function should print `[1, 4, 9, 16, 25, 36, 49, 64, 81, 100]` for $n=10$. Check that it does:

<!-- grade_id: correct_squares, points: 1, grade: true, solution: false -->
```python
"""Check that the get_collinear function returns the correct output for several all input files"""
assert get_collinear("data.txt") == {((-2, -2), (1, 1), (2, 2), (3, 3)), ((0, -3), (0, 1), (0, 5)), ((0, 1), (3, 4), (5, 6))}

```

<!-- grade_id: squares_invalid_input, points: 1, grade: true, solution: false -->
```python
"""Check that the get_collinear function raises an error for invalid input files"""
try:
    get_collinear("data.txt")
except ValueError:
    pass
else:
    raise AssertionError("did not raise")
```