## CSE 255 Programming Assignment 1: Collinear points

For this programming assignment, we'll be using a Jupyter notebook.

### In this programming assignment, you will write python3 code using pyspark to find sets of collinear points given arbitrary number of 2D points.

Definition of collinearity[1]: In geometry, collinearity of a set of points is the property of their lying on a single line. A set of points with this property is said to be collinear.

![](non-collinear-points.jpg)

Here, points P,Q,R and A,R,B are collinear. However, points A,B,C are non-collinear. For more, refer [2].

Given an input file with a set of co-ordinates, your task is to use pyspark library functions to write a program to find if three or more points are collinear.

For instance, if given these points: {(1,1), (0,1), (2,2), (3,3), (0,5), (3,4), (5,6), (0,-3), (-2,-2)}

Sets of collinear points are: {((-2,-2), (1,1), (2,2), (3,3)), ((0,1), (3,4), (5,6)), ((0,-3), (0,1), (0,5))}. Note that the ordering of the points in a set or the order of the sets does not matter. 

Note: Every set of collinear points has to have at least three points (any pair of points lie on a line).

1. https://en.wikipedia.org/wiki/Collinearity
2. http://www.mathcaptain.com/geometry/collinear-points.html

### parametrizing lines
In order to determine whether a set of points all lie on the same line we need a standard way to define (or parametrize) a line.

* One way of defining a line is as the set of points $(x,y)$ such that $y=ax+b$ for some fixed real values $a,b$.
* We call $a$ the **slope** of the line and $b$ is the $y$-intercept which is defined as the value of $y$ when $x=0$.
* This parameterization works for *almost* all lines. It does not work for vertical lines. For those lines we define $a$ to be **infinity** and $b$ to be the $x$ intercept of the line (the line is parallel to the $y$ axis so it does not intercept the $y$ axis (other than if it is the vertical line going through the origin).

To summarize, given two different points $(x_1,y_1) \neq (x_2,y_2)$, we define the parameterization $(a,b)$ as:
* **if $x_1=x_2$: ** $(\mbox{Inf},x_1)$ 
* **Else:** $(a,b)$ such that $y_1=a x_1 +b$ and $y_2=a x_2 +b$.


### Initialize spark context using 4 local cores as workers.
Note that we can create a SparkConf() object and use it to initialize the spark context. 

```python
from pyspark import SparkContext, SparkConf
conf = SparkConf().setAppName("Collinear Points").setMaster("local[4]")
sc = SparkContext(conf=conf)

from pyspark.rdd import RDD
```

### Description of the task

The goal of this assignment is to make you familiar with programming using pyspark. There are many ways to find sets of collinear points from a list of points. For the purposes of this assignment, we shall stick with the below approach:

1. List all pairs of points. You can do that efficiently in spark by computing cartesian product of the list of points with itself. For example, given three points $[(1,0), (2,0), (3,0)]$, we construct a list of nine pairs  
$[((1,0),(1,0)),((1,0), (2,0)),((1,0),(3,0))$  
$((2,0),(1,0)),((2,0), (2,0)),((2,0),(3,0))$  
$((3,0),(1,0)),((3,0), (2,0)),((3,0),(3,0))]$  

2. Remove the pairs in which the same point appears twice such as $((2,0),(2,0))$. Also, note that each pair of points appears twice, for example $((1,0),(3,0))$ and $((3,0),(1,0)$. Remove one of these pairs. **Hint:** a good way to do that is to define an ordering of points (for example: compare the $x$ values and if those are equal: compare the $y$ values. After these elimination you end up (for this example) with a list of just three pairs:  
$[((3,0),(1,0)),((3,0), (2,0)),((1,0),(2,0))]$

2. Convince yourself that if you start with $n$ different points, you end up with $\frac{n(n-1)}{2}$ pairs.

2. For each pair of points, find the parameterization $(a,b)$ of the line connecting them as described above.

3. Group the pairs according to their parameters. Clearly, if two pairs have the same $(a,b)$ values, all points in the two pairs lie on the same line.

3. Eliminate the groups that contain only one pair (any pair of points defines a line).
4. In each of the remaining groups, unpack the point-pairs to identify the individual points.
Note that if a set of points $(x_1,y_1),\ldots,(x_k,y_k)$ lie on the same line then each point will appear $k-1$ in the list of point-pairs. You therefor need to transform the list of points into sets to remove duplicates.

5. Output the sets of 3 or more colinear points.

Your task is to implement the described algorithm in Spark. You should use RDD's all the way through and collect the results into the driver only at the end.

### Below are some helper functions that you are encouraged to use in your implementations. Do not change these functions

```python
def format_result(x):
    """
    input: ((A,slope), [C1,..., Ck]) where each of A, C1,..., Ck is a point of form (Ax, Ay) and slope is a float.
    output: (C1,..., Ck, A)
    
    Concatenates collinear points.
    """
    x[1].append(x[0][0])
    return tuple(x[1])

def to_sorted_points(x):
    """
    Sorts and returns a tuple of points for further processing.
    """
    return tuple(sorted(x))
```

### Below are some functions that you will implement. You should follow the function definitions, and use them appropriately elsewhere in the notebook.

<!-- grade_id: collinear-preprocess, points: 0, grade: false, solution: true -->
```python
## Insert your answer in this cell

def get_cartesian(rdd):
    """
    Does a Cartesian product of an RDD with itself and returns an RDD with DISTINCT pairs of points.
    
    Refer:  http://spark.apache.org/docs/latest/api/python/pyspark.html?highlight=cartesian#pyspark.RDD.cartesian
    """
    
    ### BEGIN SOLUTION
    rdd = rdd.cartesian(rdd)
    
    # To remove duplicates points, example ((1,0), (1,0))
    rdd = rdd.filter(non_duplicates)
    
    return rdd
    ### END SOLUTION
```

<!-- grade_id: cell-7dfc39335c0bb589, points: 1, grade: true, solution: false -->
```python
## Let's test the to_tuple(x) implementation (1 Point)
assert to_tuple('1 1') == (1,1)
assert type(to_tuple('1 1')) == tuple
assert type(to_tuple('1 1')[0])==int and type(to_tuple('1 1')[1])==int
```

<!-- grade_id: cell-c7ddc8e33dba8d38, points: 1, grade: true, solution: false -->
```python
## Let's test the non_duplicates(x) implementation (1 Point)
assert non_duplicates(((0,0),(1,2))) == True
assert non_duplicates(((0,0),(0,0))) == False
```

<!-- grade_id: cell-6a3b8ea4640d3054, points: 3, grade: true, solution: false -->
```python
## Let's test the get_cartesian(x) implementation (3 Points)

test_rdd = sc.parallelize([(1,0), (2,0), (3,0)])
assert get_cartesian(test_rdd).collect() == [((1, 0), (2, 0)), ((1, 0), (3, 0)), ((2, 0), (1, 0)), ((2, 0), (3, 0)), ((3, 0), (1, 0)), ((3, 0), (2, 0))]
assert isinstance(get_cartesian(test_rdd), RDD) == True


### BEGIN HIDDEN TESTS
test_rdd1 = sc.parallelize([(1,0), (1,0), (1,0)])
assert get_cartesian(test_rdd1).collect() == []
assert isinstance(get_cartesian(test_rdd1), RDD) == True

test_rdd2 = sc.parallelize([(1,0)])
assert get_cartesian(test_rdd2).collect() == []
assert isinstance(get_cartesian(test_rdd2), RDD) == True

##Yoav: As much as possible: have a return messages to each assert to explain to the student what is wrong.

test_rdd3 = sc.parallelize([(1,0),(0,1)])
assert isinstance(get_cartesian(test_rdd3), RDD) == True, "get_cartesian should return an RDD"
assert get_cartesian(test_rdd3).collect() == [((1, 0), (0, 1)), ((0, 1), (1, 0))]
### END HIDDEN TESTS
```

<!-- grade_id: Collinear-slope, points: 0, grade: false, solution: true -->
```python
## Insert your answer in this cell

def find_slope(x):
    """
    Computes slope between points A and B.
    
    input: Pair (A,B) where A and B are of form (Ax, Ay) and (Bx, By) respectively. Example input: ((1,2),(3,4))
    output: Pair ((A,slope), B). Where, A and B have the same definition as input, 
            and slope refers to the slope of the line segment connecting point A and B. Example output: (((1,2),1),(3,4))
    
    NOTE: If Ax == Bx, use slope as "inf".
    """
    ### BEGIN SOLUTION
    if x[1][0] - x[0][0] == 0:
        slope = "inf"
    else:
        slope = (x[1][1] - x[0][1])/(x[1][0] - x[0][0])
    return ((x[0], slope), (x[1]))
    ### END SOLUTION



def find_collinear(rdd):
    
    """
    Finds collinear points.
    
    Input is an RDD (which is the output of the get_cartesian() function. 
    
    Output is an RDD which does the below operations:
    
    1. Find the slope of the line between all pairs of points A = (Ax, Ay) and B = (Bx, By).
    2. For each (A, B), find all points C = ((C1x, C1y), (C2x, C2y), ... (Cnx, Cny)) 
       where slope of (A,B) = slope of (A, Ci).
    3. Return (A, B, Ck) where Ck = all points of C which satisfy the condition 1.
    
    For an example, refer the assert statement test for this function.
    
    HINT:   You should use the above helper functions in conjunction with Spark RDD API (refer http://spark.apache.org/docs/latest/api/python/pyspark.html?highlight=rdd#pyspark.RDD)
            Finally, use helper function format_result() appropriately from inside this function after you have implemented the above operations.
    """
    ### BEGIN SOLUTION
    rdd = rdd.map(find_slope).groupByKey().mapValues(list).filter(lambda x: len(x[1]) > 1)
    
    rdd = rdd.map(format_result)
    
    return rdd
    ### END SOLUTION
```

<!-- grade_id: cell-1bce370d2316341e, points: 3, grade: true, solution: false -->
```python
## Let's test the find_slope(x) implementation (3 Points)

assert find_slope(((1,2),(3,4))) == (((1,2),1),(3,4))
assert find_slope(((1,2),(1,5))) == (((1,2),"inf"),(1,5))
assert find_slope(((1,2),(2,5))) == (((1,2),3),(2,5))
assert find_slope(((1,2),(-7,-2)))[0][1] == 0.5

### BEGIN HIDDEN TESTS
assert find_slope(((1,2),(1,2))) == (((1,2),"inf"),(1,2))
assert find_slope(((1,2),(0,5))) == (((1,2),-3),(0,5))
assert find_slope(((1,2),(101,5))) == (((1,2),0.03),(101,5))
### END HIDDEN TESTS
```

<!-- grade_id: cell-1506984d3ea630f3, points: 5, grade: true, solution: false -->
```python
## Let's test the find_collinear(rdd) implementation (5 Points)

test_rdd = sc.parallelize([((4, 2), (2, 1)), ((4, 2), (-3, 4)), ((4, 2), (6, 3)), ((2, 1), (4, 2)), ((2, 1), (-3, 4)), ((2, 1), (6, 3)), ((-3, 4), (4, 2)), ((-3, 4), (2, 1)), ((-3, 4), (6, 3)), ((6, 3), (4, 2)), ((6, 3), (2, 1)), ((6, 3), (-3, 4))])
assert find_collinear(test_rdd).collect() == [((4, 2), (6, 3), (2, 1)), ((2, 1), (6, 3), (4, 2)), ((4, 2), (2, 1), (6, 3))]
assert isinstance(find_collinear(test_rdd), RDD) == True


## BEGIN HIDDEN TESTS
test_rdd1 = sc.parallelize([((4, -2), (2, -1)), ((4, -2), (-3, 4)), ((2, -1), (4, -2)), ((2, -1), (-3, 4)), ((-3, 4), (4, -2)), ((-3, 4), (2, -1)), ((4, -2), (6, 3)), ((4, -2), (-9, 4)), ((4, -2), (6, -3)), ((2, -1), (6, 3)), ((2, -1), (-9, 4)), ((2, -1), (6, -3)), ((-3, 4), (6, 3)), ((-3, 4), (-9, 4)), ((-3, 4), (6, -3)), ((4, -2), (8, -4)), ((4, -2), (6, 9)), ((4, -2), (0, 0)), ((2, -1), (8, -4)), ((2, -1), (6, 9)), ((2, -1), (0, 0)), ((-3, 4), (8, -4)), ((-3, 4), (6, 9)), ((-3, 4), (0, 0)), ((4, -2), (1, 2)), ((4, -2), (2, 2)), ((4, -2), (3, 3)), ((2, -1), (1, 2)), ((2, -1), (2, 2)), ((2, -1), (3, 3)), ((-3, 4), (1, 2)), ((-3, 4), (2, 2)), ((-3, 4), (3, 3)), ((4, -2), (-9, -9)), ((2, -1), (-9, -9)), ((-3, 4), (-9, -9)), ((6, 3), (4, -2)), ((6, 3), (2, -1)), ((6, 3), (-3, 4)), ((-9, 4), (4, -2)), ((-9, 4), (2, -1)), ((-9, 4), (-3, 4)), ((6, -3), (4, -2)), ((6, -3), (2, -1)), ((6, -3), (-3, 4)), ((6, 3), (-9, 4)), ((6, 3), (6, -3)), ((-9, 4), (6, 3)), ((-9, 4), (6, -3)), ((6, -3), (6, 3)), ((6, -3), (-9, 4)), ((6, 3), (8, -4)), ((6, 3), (6, 9)), ((6, 3), (0, 0)), ((-9, 4), (8, -4)), ((-9, 4), (6, 9)), ((-9, 4), (0, 0)), ((6, -3), (8, -4)), ((6, -3), (6, 9)), ((6, -3), (0, 0)), ((6, 3), (1, 2)), ((6, 3), (2, 2)), ((6, 3), (3, 3)), ((-9, 4), (1, 2)), ((-9, 4), (2, 2)), ((-9, 4), (3, 3)), ((6, -3), (1, 2)), ((6, -3), (2, 2)), ((6, -3), (3, 3)), ((6, 3), (-9, -9)), ((-9, 4), (-9, -9)), ((6, -3), (-9, -9)), ((8, -4), (4, -2)), ((8, -4), (2, -1)), ((8, -4), (-3, 4)), ((6, 9), (4, -2)), ((6, 9), (2, -1)), ((6, 9), (-3, 4)), ((0, 0), (4, -2)), ((0, 0), (2, -1)), ((0, 0), (-3, 4)), ((8, -4), (6, 3)), ((8, -4), (-9, 4)), ((8, -4), (6, -3)), ((6, 9), (6, 3)), ((6, 9), (-9, 4)), ((6, 9), (6, -3)), ((0, 0), (6, 3)), ((0, 0), (-9, 4)), ((0, 0), (6, -3)), ((8, -4), (6, 9)), ((8, -4), (0, 0)), ((6, 9), (8, -4)), ((6, 9), (0, 0)), ((0, 0), (8, -4)), ((0, 0), (6, 9)), ((8, -4), (1, 2)), ((8, -4), (2, 2)), ((8, -4), (3, 3)), ((6, 9), (1, 2)), ((6, 9), (2, 2)), ((6, 9), (3, 3)), ((0, 0), (1, 2)), ((0, 0), (2, 2)), ((0, 0), (3, 3)), ((8, -4), (-9, -9)), ((6, 9), (-9, -9)), ((0, 0), (-9, -9)), ((1, 2), (4, -2)), ((1, 2), (2, -1)), ((1, 2), (-3, 4)), ((2, 2), (4, -2)), ((2, 2), (2, -1)), ((2, 2), (-3, 4)), ((3, 3), (4, -2)), ((3, 3), (2, -1)), ((3, 3), (-3, 4)), ((-9, -9), (4, -2)), ((-9, -9), (2, -1)), ((-9, -9), (-3, 4)), ((1, 2), (6, 3)), ((1, 2), (-9, 4)), ((1, 2), (6, -3)), ((2, 2), (6, 3)), ((2, 2), (-9, 4)), ((2, 2), (6, -3)), ((3, 3), (6, 3)), ((3, 3), (-9, 4)), ((3, 3), (6, -3)), ((-9, -9), (6, 3)), ((-9, -9), (-9, 4)), ((-9, -9), (6, -3)), ((1, 2), (8, -4)), ((1, 2), (6, 9)), ((1, 2), (0, 0)), ((2, 2), (8, -4)), ((2, 2), (6, 9)), ((2, 2), (0, 0)), ((3, 3), (8, -4)), ((3, 3), (6, 9)), ((3, 3), (0, 0)), ((-9, -9), (8, -4)), ((-9, -9), (6, 9)), ((-9, -9), (0, 0)), ((1, 2), (2, 2)), ((1, 2), (3, 3)), ((2, 2), (1, 2)), ((2, 2), (3, 3)), ((3, 3), (1, 2)), ((3, 3), (2, 2)), ((1, 2), (-9, -9)), ((2, 2), (-9, -9)), ((3, 3), (-9, -9)), ((-9, -9), (1, 2)), ((-9, -9), (2, 2)), ((-9, -9), (3, 3))])
assert find_collinear(test_rdd1).collect() == [((4, -2), (6, -3), (8, -4), (0, 0), (2, -1)), ((6, -3), (6, 9), (6, 3)), ((4, -2), (2, -1), (6, -3), (0, 0), (8, -4)), ((4, -2), (2, -1), (6, -3), (8, -4), (0, 0)), ((4, -2), (2, -1), (8, -4), (0, 0), (6, -3)), ((0, 0), (2, 2), (-9, -9), (3, 3)), ((0, 0), (2, 2), (3, 3), (-9, -9)), ((2, -1), (6, -3), (8, -4), (0, 0), (4, -2)), ((6, 3), (6, 9), (6, -3)), ((6, 3), (6, -3), (6, 9)), ((2, 2), (3, 3), (-9, -9), (0, 0)), ((0, 0), (3, 3), (-9, -9), (2, 2))]
assert isinstance(find_collinear(test_rdd1), RDD) == True
# END HIDDEN TESTS
```

<!-- grade_id: cell-bfd990be780ce6a8, points: 2, grade: true, solution: false -->
```python
## Let's test find_collinear(rdd) when run with output of get_cartesian(rdd) (2 Points)

test_rdd = sc.parallelize([(4, -2), (2, -1), (-3,4), (6,3), (-9,4), (6, -3), (8,-4), (6,9)])
test_rdd = get_cartesian(test_rdd)
assert find_collinear(test_rdd).collect() == [((6, -3), (6, 9), (6, 3)), ((4, -2), (2, -1), (8, -4), (6, -3)), ((2, -1), (6, -3), (8, -4), (4, -2)), ((6, 3), (6, 9), (6, -3)), ((4, -2), (6, -3), (8, -4), (2, -1)), ((6, 3), (6, -3), (6, 9)), ((4, -2), (2, -1), (6, -3), (8, -4))]

## BEGIN HIDDEN TESTS
test_rdd1 = sc.parallelize([(4, -2), (2, -1), (-3,4), (6,3), (-9,4), (6, -3), (8,-4), (6,9), (0,0),(1,2),(2,2),(3,3),(-9,-9)])
test_rdd1 = get_cartesian(test_rdd1)
assert find_collinear(test_rdd1).collect() == [((4, -2), (2, -1), (6, -3), (8, -4), (0, 0)), ((0, 0), (3, 3), (-9, -9), (2, 2)), ((6, -3), (6, 9), (6, 3)), ((4, -2), (2, -1), (8, -4), (0, 0), (6, -3)), ((0, 0), (2, 2), (3, 3), (-9, -9)), ((2, -1), (6, -3), (8, -4), (0, 0), (4, -2)), ((6, 3), (6, 9), (6, -3)), ((2, 2), (3, 3), (-9, -9), (0, 0)), ((4, -2), (6, -3), (8, -4), (0, 0), (2, -1)), ((6, 3), (6, -3), (6, 9)), ((4, -2), (2, -1), (6, -3), (0, 0), (8, -4)), ((0, 0), (2, 2), (-9, -9), (3, 3))]
## END HIDDEN TESTS
```

### Add tests for each individual function

**Yoav:** Including checking thaat the output is an RDD with the correct format.

### This cell below runs all of the above functions step by step. Do not change this cell.

```python
def get_collinear(filename):
    """
    Computes the set of collinear points
    """
    
    # Load the data file into an RDD
    rdd = sc.textFile(filename)
    
    # Transform the loaded points into tuples
    rdd = rdd.map(to_tuple)
    
    # Transform the RDD to now have the cartesian product of the RDD with itself
    rdd = get_cartesian(rdd)
    
    # Transform the RDD to now just have sets of Collinear Points 
    rdd = find_collinear(rdd)
    
    # Sorting each of your returned sets of collinear points. This is for grading purposes. You may ignore this.
    rdd = rdd.map(to_sorted_points)
    
    # Collecting the collinear points RDD in a set to remove duplicate sets of collinear points. This is for grading purposes. You may ignore this.
    res = set(rdd.collect())
    
    return res
```

<!-- grade_id: Collinear_soln_visible, points: 5, grade: true, solution: false -->
```python
"""
    Check that the get_collinear() function returns the correct output for several input files.
    You may assume that input files do not have duplicate points. (10 points)
    
"""
assert get_collinear("data.txt") == {((-2, -2), (1, 1), (2, 2), (3, 3)), ((0, 1), (3, 4), (5, 6)), ((0, -3), (0, 1), (0, 5))}
assert get_collinear("data50.txt") == {((3, 6), (7, 4), (9, 3)), ((1, 6), (3, 6), (4, 6), (7, 6)), 
                                 ((0, 2), (3, 1), (6, 0)), ((1, 0), (2, 0), (5, 0), (6, 0)), 
                                 ((1, 3), (3, 6), (5, 9)), ((0, 8), (4, 6), (6, 5)), 
                                 ((6, 0), (6, 1), (6, 5), (6, 9)), 
                                 ((7, 2), (7, 3), (7, 4), (7, 6), (7, 8)), ((3, 1), (3, 3), (3, 6)), 
                                 ((0, 2), (1, 2), (5, 2), (7, 2)), ((0, 3), (2, 5), (3, 6), (6, 9)), 
                                 ((0, 2), (1, 3), (2, 4), (4, 6), (5, 7)), ((1, 2), (4, 3), (7, 4)), 
                                 ((0, 3), (4, 6), (8, 9)), ((9, 3), (9, 4), (9, 5)), ((2, 5), (5, 7), (8, 9)), 
                                 ((0, 5), (2, 4), (4, 3), (8, 1)), ((0, 8), (1, 6), (2, 4)), 
                                 ((3, 6), (5, 2), (6, 0)), ((5, 9), (6, 9), (8, 9)), 
                                 ((0, 8), (1, 8), (7, 8)), ((0, 4), (1, 3), (3, 1)), ((5, 9), (7, 6), (9, 3)), 
                                 ((1, 2), (2, 4), (3, 6)), ((0, 7), (1, 5), (3, 1)), 
                                 ((1, 5), (2, 4), (3, 3), (6, 0)), ((0, 2), (3, 3), (9, 5)), 
                                 ((0, 7), (1, 6), (2, 5), (4, 3), (5, 2), (6, 1)), 
                                 ((0, 4), (1, 5), (5, 9)), ((1, 5), (3, 6), (5, 7), (7, 8)), 
                                 ((1, 6), (3, 3), (5, 0)), ((3, 6), (4, 3), (5, 0)), 
                                 ((1, 2), (4, 5), (7, 8), (8, 9)), ((0, 2), (1, 1), (2, 0)), 
                                 ((3, 3), (4, 5), (5, 7), (6, 9)), ((0, 2), (0, 3), (0, 4), (0, 5), (0, 7), (0, 8)), 
                                 ((2, 0), (4, 3), (8, 9)), ((5, 7), (6, 5), (7, 3), (8, 1)), ((5, 0), (7, 6), (8, 9)), 
                                 ((5, 0), (6, 1), (7, 2), (9, 4)), ((0, 4), (1, 2), (2, 0)), 
                                 ((1, 1), (3, 1), (6, 1), (8, 1)), ((5, 7), (7, 6), (9, 5)), ((1, 1), (7, 4), (9, 5)), 
                                 ((0, 4), (2, 4), (7, 4), (9, 4)), ((1, 0), (3, 1), (5, 2), (7, 3), (9, 4)), 
                                 ((2, 0), (3, 3), (4, 6), (5, 9)), ((4, 3), (4, 5), (4, 6)), 
                                 ((1, 0), (4, 3), (6, 5), (7, 6)), ((0, 3), (2, 4), (4, 5)), 
                                 ((1, 6), (4, 5), (7, 4)), ((1, 0), (1, 1), (1, 2), (1, 3), (1, 5), (1, 6), (1, 8)), 
                                 ((0, 3), (1, 3), (3, 3), (4, 3), (7, 3), (9, 3)), ((0, 4), (2, 5), (4, 6)), 
                                 ((0, 7), (3, 6), (6, 5), (9, 4)), ((1, 8), (4, 6), (7, 4)), 
                                 ((0, 5), (3, 3), (6, 1)), ((1, 8), (3, 6), (4, 5), (7, 2), (8, 1)), 
                                 ((1, 2), (3, 1), (5, 0)), ((1, 1), (5, 2), (9, 3)), 
                                 ((5, 0), (5, 2), (5, 7), (5, 9)), ((0, 5), (1, 5), (2, 5), (4, 5), (6, 5), (9, 5)), 
                                 ((3, 1), (4, 5), (5, 9)), ((2, 0), (2, 4), (2, 5)), ((5, 2), (6, 5), (7, 8))}
# assert execute("test.txt") == {((1, 0), (2, 0), (3, 0))}
```

<!-- grade_id: Collinear_soln_hidden1, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
import ast

with open('data100_res.txt','r') as f:
     d = ast.literal_eval(f.read())
assert execute("data100.txt") == d
### END HIDDEN TESTS
```

<!-- grade_id: Collinear_soln_hidden2, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
with open('data200_res.txt','r') as f:
     d = ast.literal_eval(f.read())
assert execute("data200.txt") == d
### END HIDDEN TESTS
```

<!-- grade_id: Collinear_soln_hidden3, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
with open('data400_res.txt','r') as f:
     d = ast.literal_eval(f.read())
assert execute("data400.txt") == d
### END HIDDEN TESTS
```

<!-- grade_id: Collinear_soln_hidden4, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
with open('data800_res.txt','r') as f:
     d = ast.literal_eval(f.read())
assert get_collinear("data800.txt") == d
### END HIDDEN TESTS
```

<!-- grade_id: Collinear_invalid_input, points: 5, grade: true, solution: false -->
```python
## TODO: Not sure we need this cell in the released HW. If not, please delete it.

## 0 points
"""Check that the get_collinear() function raises an error for invalid input files"""
try:
    get_collinear("data.txt")
except ValueError:
    raise AssertionError("Data points do not match the expected format.")
else:
    pass
```