<problem>
  <text>
    <script type="loncapa/python">
# random variables (no need to import random library)
# import random
n1 = random.randrange(2,4,1)
n2 = random.randrange(5,7,1)
pda = random.randrange(10,40,1)/1000.
pdb = random.randrange(20,50,1)/1000.
pdc = random.randrange(30,60,1)/1000.

fac = ["A","B","C"]
ind = random.randrange(0,3,1)
refac = fac[ind]
# print ind,refac
num = ["{0}*{1}".format(n1,pda), "{0}*{1}".format(n2,pdb), "{0}".format(pdc)]
enum = num[ind]
denom = "{0}*{1}+{2}*{3}+{4}".format(n1, pda, n2, pdb, pdc)

# Solutions with variables converted to string
# Make sure you name the solution with part id at the end. e.g. 'solution1' will be solution for part 1. 
solution1 = "{0}/({1})".format(enum,denom)

# Group all solutions into a list
solutions = [solution1]
# solutions


# hint text id
hint_text_id = ["cse103fall2016week1problem1part1","cse103fall2016week1problem1part2"]
    </script>

    <script type="text/javascript">
      //fill in number of hint
        var hint_number = 2;

      //fill in hint input box id
        var hintIds = ["input_57ccdc78397649aca59b419c95573254_2_1",
                "input_1ddc9e0c3eea4ba3b61b291a7cfb650d_2_1"];

      //fill in problem input box id
        var proIds = ["input_0463a3ce78dd46429e9e2fcf681616d9_2_1",
               "input_0463a3ce78dd46429e9e2fcf681616d9_3_1"];
               
      //fill in hint text id. Need to match the hint_text_id from above
        var hintTextIds = ["cse103fall2016week1problem1part1",
                  "cse103fall2016week1problem1part2"];
    </script>



<p>Factories A, B and C produce computers.</p>
<ul>
<li>
<p>Factory A produces $n1 times as many computers as factory C. And factory B produces $n2 times as many computers as factory C.</p>
</li>
<li>
<p>The probabilities that each factory produces a defective computer are</p>
</li>
<li>For factory A, the probability is $pda</li>
<li>For factory B, the probability is $pdb</li>
<li>For factory C, the probability is $pdc</li>
</ul>
<p>A computer is selected at random and it is found to be defective.
What is the probability it came from factory $refac?</p>


    <customresponse cfn="check" expect="\[$solution1\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


    

    

    <script type="loncapa/python">
from hint import evaluate
from hint.get_hint import get_hint

# grader of the problem
def check(expect, ans):
  return evaluate.evaluate(expect, ans)

# hint function
def hint_fn(answer_ids, student_answers, new_cmap, old_cmap):
  aid = answer_ids[0]
  ans = str(student_answers[aid])
  i = int(aid[-3])-2

  if ans == "":
    return
  key = new_cmap.keys()[0]
  if new_cmap[key]['correctness'] == "correct":
    return

  hint = get_hint(solutions[i], ans, hint_text_id, i)
  new_cmap.set_hint_and_mode(aid,hint,'always')
    </script>

        
    <script type="text/javascript" src="/static/js/adaptive_hint.js"></script>
    <script type="text/javascript">                             
      for (var i=0; i != hint_number; i++) {
        hide_hint(i);
      }
      show_hint(hint_number);
    </script>

  </text>
</problem>