<problem>
  <text>
    <script type="loncapa/python">
variable_values = {'mean':range(-5,5), 'std':range(1,11),'n':range(10,110,10)}
print variable_values
mu = random.randrange(-5,5,1)
std = random.randrange(1,11,1)
n = random.randrange(10,110,10)
solution1= "{}".format(mu)
solution2="({0}^2)/{1}".format(std, n)
solution3='(1/4)^2'
solution4='2*Q(4)'
solutions = [solution1, solution2, solution3, solution4]
solutions

# hint text id
hint_text_id = ["cse103fall2016week1problem1part1","cse103fall2016week1problem1part2"]
    </script>

    <script type="text/javascript">
      //fill in number of hint
        var hint_number = 2;

      //fill in hint input box id
        var hintIds = ["input_57ccdc78397649aca59b419c95573254_2_1",
                "input_1ddc9e0c3eea4ba3b61b291a7cfb650d_2_1"];

      //fill in problem input box id
        var proIds = ["input_0463a3ce78dd46429e9e2fcf681616d9_2_1",
               "input_0463a3ce78dd46429e9e2fcf681616d9_3_1"];
               
      //fill in hint text id. Need to match the hint_text_id from above
        var hintTextIds = ["cse103fall2016week1problem1part1",
                  "cse103fall2016week1problem1part2"];
    </script>



<h2>Chebyshev and CLT</h2>
<p>Suppose \(X_1,X_2,\ldots X_n\) are IID non-negative random variables with common mean \(E(X_i)=$mu\) and common variance \(\text{std}^2=$std ^2 = Var(X_i)\).</p>
<p>Suppose \(S=\frac{1}{n} \sum_{i=1}^n X_i\)  is the average of the \(X_i\)s.</p>
<p>Let \( n = $n \).</p>
<p>1] What is \(E(S)=\)?</p>


    <customresponse cfn="check" expect="\[$solution1\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<p>2] What is \(Var(S)\)</p>


    <customresponse cfn="check" expect="\[$solution2\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<p>3] Using Chernoff, what is the bound you can put on \(P(|X-\mu|&gt;4\sigma)\)</p>


    <customresponse cfn="check" expect="\[$solution3\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<p>4] Assuming that \(n\) is sufficiently large that Central Limit theorem applies on \(S\), what is the bound that you can put on \(P(|X-\mu|&gt;4\sigma)\)</p>


    <customresponse cfn="check" expect="\[$solution4\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


    

    

    <script type="loncapa/python">
from hint import evaluate
from hint.get_hint import get_hint

# grader of the problem
def check(expect, ans):
  return evaluate.evaluate(expect, ans)

# hint function
def hint_fn(answer_ids, student_answers, new_cmap, old_cmap):
  aid = answer_ids[0]
  ans = str(student_answers[aid])
  i = int(aid[-3])-2

  if ans == "":
    return
  key = new_cmap.keys()[0]
  if new_cmap[key]['correctness'] == "correct":
    return

  hint = get_hint(solutions[i], ans, hint_text_id, i)
  new_cmap.set_hint_and_mode(aid,hint,'always')
    </script>

        
    <script type="text/javascript" src="/static/js/adaptive_hint.js"></script>
    <script type="text/javascript">                             
      for (var i=0; i != hint_number; i++) {
        hide_hint(i);
      }
      show_hint(hint_number);
    </script>

  </text>
</problem>