<problem>
  <text>
    <script type="loncapa/python">

k = random.randrange(12,16,1)

r = random.randrange(1,2,1)
a = random.randrange(1,2,1)
select = random.randrange(8,15,1)

solution1 = "C({0} + 3,3)".format(k)
solution2 = "C({0} + 2,2)".format(k)
solution3 = "C({0} - 1 + 2,2)".format(k)
solution4 = "C({0} - 2 + 2,2)".format(k)
solution5 = "C({0} - 2 + 2,2) + C({0} - 1 + 2,2) + C({0} + 2,2)".format(k)

solutions = [solution1, solution2, solution3, solution4, solution5]


# hint text id
hint_text_id = ["cse103fall2016week1problem1part1","cse103fall2016week1problem1part2"]
    </script>

    <script type="text/javascript">
      //fill in number of hint
        var hint_number = 2;

      //fill in hint input box id
        var hintIds = ["input_57ccdc78397649aca59b419c95573254_2_1",
                "input_1ddc9e0c3eea4ba3b61b291a7cfb650d_2_1"];

      //fill in problem input box id
        var proIds = ["input_0463a3ce78dd46429e9e2fcf681616d9_2_1",
               "input_0463a3ce78dd46429e9e2fcf681616d9_3_1"];
               
      //fill in hint text id. Need to match the hint_text_id from above
        var hintTextIds = ["cse103fall2016week1problem1part1",
                  "cse103fall2016week1problem1part2"];
    </script>



<h4>Distributing Sweets</h4>
<p>We are trying to distribute $k sweets to 4 children: James, Apple, Sally, and Mary.</p>
<p>How many ways are there to distribute the sweets?</p>


    <customresponse cfn="check" expect="\[$solution1\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<p>How many ways are there to distribute the sweets so that James gets no candy?</p>


    <customresponse cfn="check" expect="\[$solution2\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<p>How many ways are there to distribute the sweets so that James gets exactly 1 candy?</p>


    <customresponse cfn="check" expect="\[$solution3\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<p>How many ways are there to distribute the sweets so that James gets exactly 2 candies?</p>


    <customresponse cfn="check" expect="\[$solution4\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<p>How many ways are there to distribute the sweets so that James gets at most 2 candies?</p>


    <customresponse cfn="check" expect="\[$solution5\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


    

    

    <script type="loncapa/python">
from hint import evaluate
from hint.get_hint import get_hint

# grader of the problem
def check(expect, ans):
  return evaluate.evaluate(expect, ans)

# hint function
def hint_fn(answer_ids, student_answers, new_cmap, old_cmap):
  aid = answer_ids[0]
  ans = str(student_answers[aid])
  i = int(aid[-3])-2

  if ans == "":
    return
  key = new_cmap.keys()[0]
  if new_cmap[key]['correctness'] == "correct":
    return

  hint = get_hint(solutions[i], ans, hint_text_id, i)
  new_cmap.set_hint_and_mode(aid,hint,'always')
    </script>

        
    <script type="text/javascript" src="/static/js/adaptive_hint.js"></script>
    <script type="text/javascript">                             
      for (var i=0; i != hint_number; i++) {
        hide_hint(i);
      }
      show_hint(hint_number);
    </script>

  </text>
</problem>