<problem>
  <text>
    <script type="loncapa/python">
import random
pF=random.randrange(5,30,1)
pEF=random.randrange(1,pF,1)
pF/=100.
pEF/=100.
print pF,pEF
solution1 = '{0}/{1}'.format(pEF,pF)
solution2 = '0.5-{}'.format(pEF)
solution3 = '1-{}'.format(pF)
solution4 = '({0})/({1})'.format(solution2,solution3)

solutions = [solution1, solution2,solution3, solution4]
solutions

# hint text id
hint_text_id = ["cse103fall2016week1problem1part1","cse103fall2016week1problem1part2"]
    </script>

    <script type="text/javascript">
      //fill in number of hint
        var hint_number = 2;

      //fill in hint input box id
        var hintIds = ["input_57ccdc78397649aca59b419c95573254_2_1",
                "input_1ddc9e0c3eea4ba3b61b291a7cfb650d_2_1"];

      //fill in problem input box id
        var proIds = ["input_0463a3ce78dd46429e9e2fcf681616d9_2_1",
               "input_0463a3ce78dd46429e9e2fcf681616d9_3_1"];
               
      //fill in hint text id. Need to match the hint_text_id from above
        var hintTextIds = ["cse103fall2016week1problem1part1",
                  "cse103fall2016week1problem1part2"];
    </script>



<h2>Bayesian Inference, party affiliation and Spicy Food</h2>
<p>You meet a stranger, a random citizen of the United States. What's the chance that he or she votes for the Democratic party?</p>
<p>\[ \Omega = {\mbox{citizens of the US}}, \ \ E = {\omega \in \Omega: \omega\; \mbox{votes Dem}}.\]</p>
<p>Probably the probability is close to 50%, lets assume assume \(P(E) = 0.5\)</p>
<p>suppose you find out some additional information - whether of not the person likes spicy food. Or in other words, whether the person is in the set \(F\):</p>
<p>\[ F = {\omega \in \Omega: \omega\; \mbox{likes spicy food}}.\]</p>
<p>Suppose you know the following probabilities:</p>
<p>\[ P(F) = $pF, P(E \cap F) = $pEF \]</p>
<p>1) Suppose that the person <strong>likes</strong> spicy food, How does that change the probability that the person is a democrat? In other words, what is \(P(E|F)\)?</p>


    <customresponse cfn="check" expect="\[$solution1\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<p>2) You know that \(P(E \cap F) = $pEF \) What is \(P(E \cap F^c)\) (the probability that a random person is a democrat that does not like spicy food).</p>


    <customresponse cfn="check" expect="\[$solution2\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<p>3) What is the probability of people that do not like spicy food \(P(F^c)=\)</p>


    <customresponse cfn="check" expect="\[$solution3\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<p>4) Suppose that the person <strong>dislikes</strong> spicy food, how does that change the probability that the person is a democrat? In other words, what is \(P(E|F^c)\)?</p>


    <customresponse cfn="check" expect="\[$solution4\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


    

    

    <script type="loncapa/python">
from hint import evaluate
from hint.get_hint import get_hint

# grader of the problem
def check(expect, ans):
  return evaluate.evaluate(expect, ans)

# hint function
def hint_fn(answer_ids, student_answers, new_cmap, old_cmap):
  aid = answer_ids[0]
  ans = str(student_answers[aid])
  i = int(aid[-3])-2

  if ans == "":
    return
  key = new_cmap.keys()[0]
  if new_cmap[key]['correctness'] == "correct":
    return

  hint = get_hint(solutions[i], ans, hint_text_id, i)
  new_cmap.set_hint_and_mode(aid,hint,'always')
    </script>

        
    <script type="text/javascript" src="/static/js/adaptive_hint.js"></script>
    <script type="text/javascript">                             
      for (var i=0; i != hint_number; i++) {
        hide_hint(i);
      }
      show_hint(hint_number);
    </script>

  </text>
</problem>