<problem>
  <text>
    <script type="loncapa/python">

# random variables (no need to import random library)
import math

ans1 = 0
mu = 0 
a = 0
x = 0
lambd = 0
while (0.002 > ans1) :
    mu = random.randrange(50,70,5) / 10.
    lambd = 1. / mu
    a = random.randrange(9,12,1)
    x = lambd * a
    ans1 = math.exp(x)
    k=random.randrange(3,7,1)

# Solutions with variables converted to string
# Make sure you name the solution with part id at the end. e.g. 'solution1' will be solution for part 1.
solution1 = "e^{{(-1/{0} * {1})}}".format(mu,a)
solution2 = "e^{{(-1/{0}*10)}}*((-1/{0}*10)^{{{1}}})/({1}!)".format(mu,k)

# Group all solutions into a list
solutions = [solution1,solution2]

# hint text id
hint_text_id = ["cse103fall2016week1problem1part1","cse103fall2016week1problem1part2"]
    </script>

    <script type="text/javascript">
      //fill in number of hint
        var hint_number = 2;

      //fill in hint input box id
        var hintIds = ["input_57ccdc78397649aca59b419c95573254_2_1",
                "input_1ddc9e0c3eea4ba3b61b291a7cfb650d_2_1"];

      //fill in problem input box id
        var proIds = ["input_0463a3ce78dd46429e9e2fcf681616d9_2_1",
               "input_0463a3ce78dd46429e9e2fcf681616d9_3_1"];
               
      //fill in hint text id. Need to match the hint_text_id from above
        var hintTextIds = ["cse103fall2016week1problem1part1",
                  "cse103fall2016week1problem1part2"];
    </script>



<h3>Time serviced by the  Cashier</h3>
<p>The manager of a supermarket tracked the amount of time needed for customers to be served by the cashier.  After checking with his statistics professor, he concluded that the checkout times are exponentially distributed with a mean of \($mu\) minutes.  What propotion of customers require more than \($a\) minutes to check out?</p>
<ul>
<li>Proportion =</li>
</ul>


    <customresponse cfn="check" expect="\[$solution1\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<ul>
<li>What is the probability that exactly \($k\) customers were serviced between 10am and 10:10am ? (Hint, this number has a Poisson distribution).</li>
</ul>
<p>(we count customers whose service <strong>started</strong> between 10am and 10:10am but not customers whose service period started before 10am and ended after 10am)</p>


    <customresponse cfn="check" expect="\[$solution2\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>



    
    

    

    <script type="loncapa/python">
from hint import evaluate
from hint.get_hint import get_hint

# grader of the problem
def check(expect, ans):
  return evaluate.evaluate(expect, ans)

# hint function
def hint_fn(answer_ids, student_answers, new_cmap, old_cmap):
  aid = answer_ids[0]
  ans = str(student_answers[aid])
  i = int(aid[-3])-2

  if ans == "":
    return
  key = new_cmap.keys()[0]
  if new_cmap[key]['correctness'] == "correct":
    return

  hint = get_hint(solutions[i], ans, hint_text_id, i)
  new_cmap.set_hint_and_mode(aid,hint,'always')
    </script>

        
    <script type="text/javascript" src="/static/js/adaptive_hint.js"></script>
    <script type="text/javascript">                             
      for (var i=0; i != hint_number; i++) {
        hide_hint(i);
      }
      show_hint(hint_number);
    </script>

  </text>
</problem>