<problem>
  <text>
    <script type="loncapa/python">

R = random.randrange(4,9,1)
RL = R - 1

solution1 = "{0}^4".format(R)
solution2 = "4"
solution3 = "({0}-1)*{0}".format(R)
solution4 = "{0}*({0}-1)/({0}^4)".format(R)

solutions = [solution1,solution2,solution3,solution4]
# hint text id
hint_text_id = ["cse103fall2016week1problem1part1","cse103fall2016week1problem1part2"]
    </script>

    <script type="text/javascript">
      //fill in number of hint
        var hint_number = 2;

      //fill in hint input box id
        var hintIds = ["input_57ccdc78397649aca59b419c95573254_2_1",
                "input_1ddc9e0c3eea4ba3b61b291a7cfb650d_2_1"];

      //fill in problem input box id
        var proIds = ["input_0463a3ce78dd46429e9e2fcf681616d9_2_1",
               "input_0463a3ce78dd46429e9e2fcf681616d9_3_1"];
               
      //fill in hint text id. Need to match the hint_text_id from above
        var hintTextIds = ["cse103fall2016week1problem1part1",
                  "cse103fall2016week1problem1part2"];
    </script>



<p>Suppose a seqeunce of 4 digits in the range 1-$R is chosen uniformly at random. What is the probability that the first and third digit areequal, the second and the fourth digits are equal, but the first and second digits are unequal?</p><ul><li><p>We are going to use the standard formula for probability of sets in a discrete uniform distribution space: the probability of a set \(A\) in an event space \(\Omega\) is</p><p>\[P(A) = \frac{|A|}{|\Omega|}\]</p></li><li><p>In the denominator we place the <em>size</em> of the sample space. This is easy, as we have $R choices in each of the four locations the number of sequences, which is equal to the size of the sample space, thus \(|\Omega| = \)</p></li></ul>

    <customresponse cfn="check" expect="\[$solution1\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<ul><li>Now we want to compute the size of the set \(A\). To do that it is useful to realize that we need only concern ourselves with the choices of two digits, say digits 1 and 2.  This is because digit 1 determines the value of digit 3. Which digit does digit 2 determines?</li></ul>

    <customresponse cfn="check" expect="\[$solution2\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<ul><li>Now, how many choices do we have for digit 1? Obviously, any integer in the range \(1,\ldots,$R\), or $Rchoices. Once the first digit has been chosen, we have one less choice for the second digit, in other words $RL. The size of the set \(A\) is the product of these two numbers:</li></ul>

    <customresponse cfn="check" expect="\[$solution3\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<ul><li>Putting this all together we get that the probability of \(A\) is</li></ul><p>\(P(A)=\)</p>

    <customresponse cfn="check" expect="\[$solution4\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>



    
    

    

    <script type="loncapa/python">
from hint import evaluate
from hint.get_hint import get_hint

# grader of the problem
def check(expect, ans):
  return evaluate.evaluate(expect, ans)

# hint function
def hint_fn(answer_ids, student_answers, new_cmap, old_cmap):
  aid = answer_ids[0]
  ans = str(student_answers[aid])
  i = int(aid[-3])-2

  if ans == "":
    return
  key = new_cmap.keys()[0]
  if new_cmap[key]['correctness'] == "correct":
    return

  hint = get_hint(solutions[i], ans, hint_text_id, i)
  new_cmap.set_hint_and_mode(aid,hint,'always')
    </script>

        
    <script type="text/javascript" src="/static/js/adaptive_hint.js"></script>
    <script type="text/javascript">                             
      for (var i=0; i != hint_number; i++) {
        hide_hint(i);
      }
      show_hint(hint_number);
    </script>

  </text>
</problem>