<problem>
  <text>
    <script type="loncapa/python">

#### Written by Yoav Freund 11/3/2016

# random variables (no need to import random library)
def rand_vals(i=1,j=100,n=2):
    return [random.randrange(i,j,1) for k in range(n)]
a_rand=rand_vals() ; b_rand=rand_vals() ; c_rand=rand_vals();
a_vals=[]; b_vals=[]; c_vals=[];
for i in range(len(a_rand)):
    for j in range(len(b_rand)):
    	for k in range(len(c_rand)):
        	a_vals.append(a_rand[i])
        	b_vals.append(b_rand[j])
        	c_vals.append(c_rand[k])
variable_values = {'a':a_vals, 'b':b_vals, 'v':c_vals}

# Random valued coefficients:
[c,d,e,f,g]=rand_vals(i=2,j=5,n=5)

# Solutions with variables converted to string
# Make sure you name the solution with part id at the end. e.g. 'solution1' will be solution for part 1. 
solution1 = "{0}*a".format(c)
solution2 = "{0}*a+{1}*b+{2}".format(c,d,e)
solution3 = "-{0}*a+{1}*b+{2}*{3}*v-1".format(c,d,c,d) 

# Group all solutions into a list
solutions = [solution1, solution2, solution3]

# hint id
hint_text_id = ["cse103fall2016week1problem1part1","cse103fall2016week1problem1part2"]
    </script>

    <script type="text/javascript">
      //fill in number of hint
        var hint_number = 2;

      //fill in hint input box id
        var hintIds = ["input_57ccdc78397649aca59b419c95573254_2_1",
                "input_1ddc9e0c3eea4ba3b61b291a7cfb650d_2_1"];

      //fill in problem input box id
        var proIds = ["input_0463a3ce78dd46429e9e2fcf681616d9_2_1",
               "input_0463a3ce78dd46429e9e2fcf681616d9_3_1"];
               
      //fill in hint text id. Need to match the hint_text_id from above
        var hintTextIds = ["cse103fall2016week1problem1part1",
                  "cse103fall2016week1problem1part2"];
    </script>



<h1>Expectation and Variance calculations</h1><p>Suppose \(X,Y\) are random variables and that \(a=E(X), b=E(Y), v=E(XY)\) </p><p>Express the following using \(a\) and \(b\)</p><p>(1) \(E($c X)\)</p>

    <customresponse cfn="check" expect="\[$solution1\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<p>(2) \(E($c * X+ $d * Y+ $e)\)</p>

    <customresponse cfn="check" expect="\[$solution2\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<p>(3) \(E(($c * X+1) * ($d * Y-1))\)</p>

    <customresponse cfn="check" expect="\[$solution3\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>



    
    

    

    <script type="loncapa/python">
from hint import evaluate
from hint.get_hint import get_hint_w_variables

# grader of the problem
def check(expect, ans):
  return evaluate.evaluate_w_variables(expect, ans, variable_values)

# hint function
def hint_fn(answer_ids, student_answers, new_cmap, old_cmap):
  aid = answer_ids[0]
  ans = str(student_answers[aid])
  i = int(aid[-3])-2

  if ans == "":
    return
  key = new_cmap.keys()[0]
  if new_cmap[key]['correctness'] == "correct":
    return

  hint = get_hint_w_variables(solutions[i], ans, hint_text_id, i, variable_values)
  new_cmap.set_hint_and_mode(aid,hint,'always')
    </script>

        
    <script type="text/javascript" src="/static/js/adaptive_hint.js"></script>
    <script type="text/javascript">                             
      for (var i=0; i != hint_number; i++) {
        hide_hint(i);
      }
      show_hint(hint_number);
    </script>

  </text>
</problem>