<problem>
  <text>
    <script type="loncapa/python">

R1=random.randrange(8,15,1);
R2=random.randrange(2,R1-2,1);

solution1 = "{0} - 1".format(R1)
solution2 = "{0}".format(R2)
solution3 = "(({1}+{0}-1)!)/(({1})!*({0}-1)!)".format(R1,R2)

solutions = [solution1, solution2, solution3]
# hint text id
hint_text_id = ["cse103fall2016week1problem1part1","cse103fall2016week1problem1part2"]
    </script>

    <script type="text/javascript">
      //fill in number of hint
        var hint_number = 2;

      //fill in hint input box id
        var hintIds = ["input_57ccdc78397649aca59b419c95573254_2_1",
                "input_1ddc9e0c3eea4ba3b61b291a7cfb650d_2_1"];

      //fill in problem input box id
        var proIds = ["input_0463a3ce78dd46429e9e2fcf681616d9_2_1",
               "input_0463a3ce78dd46429e9e2fcf681616d9_3_1"];
               
      //fill in hint text id. Need to match the hint_text_id from above
        var hintTextIds = ["cse103fall2016week1problem1part1",
                  "cse103fall2016week1problem1part2"];
    </script>



<p>You walk into a candy store and notice that there are five types of candy. Your mother allows you to pick exactly three pieces of candy, of whichever type(s) you want. How many ways are there to do this?</p><p>You can represent the outcome by 5-tuple \((n_1, n_2, \ldots, n_5)\) in which \(n_i\) is the number of pieces of the \(i\) th type of candy. How many such tuples are there, subject to \(n_1 + n_2 + \cdots + n_5 = 3\)</p><p>To answer the question, we'll represent each tuple in a different format, as a sequence of length 7 containing three stars and four bars. For instance, the sequence \(| *  * | | | * \) denotes \((0,2,0,0,1)\) (two candies of type 2 and one candy of type 5): the number of candies of type \(i\) is the number of stars between the \((i-1)\) th and \(i\) th bars.</p><p>So we have rephrased the question thus: how many sequences are there with four bars and three stars? Well, this is a sequence of size 7, and we must pick three of the seven positions at which to place stars. The number of such choices is \({7 \choose 3} = 35\)</p><p>Suppose balls come in $R1 colors and that you are to pick out $R2 balls. How many different color combinations are possible?</p><p>By the analogy explained above, how many bars do you need</p>

    <customresponse cfn="check" expect="\[$solution1\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<p>By the analogy explained above, how many stars do you need</p>

    <customresponse cfn="check" expect="\[$solution2\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<p>The number of color combinations is</p>

    <customresponse cfn="check" expect="\[$solution3\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>



    
    

    

    <script type="loncapa/python">
from hint import evaluate
from hint.get_hint import get_hint

# grader of the problem
def check(expect, ans):
  return evaluate.evaluate(expect, ans)

# hint function
def hint_fn(answer_ids, student_answers, new_cmap, old_cmap):
  aid = answer_ids[0]
  ans = str(student_answers[aid])
  i = int(aid[-3])-2

  if ans == "":
    return
  key = new_cmap.keys()[0]
  if new_cmap[key]['correctness'] == "correct":
    return

  hint = get_hint(solutions[i], ans, hint_text_id, i)
  new_cmap.set_hint_and_mode(aid,hint,'always')
    </script>

        
    <script type="text/javascript" src="/static/js/adaptive_hint.js"></script>
    <script type="text/javascript">                             
      for (var i=0; i != hint_number; i++) {
        hide_hint(i);
      }
      show_hint(hint_number);
    </script>

  </text>
</problem>