<problem>
  <text>
    <script type="loncapa/python">

a = random.randrange(300,400,1)
b = random.randrange(15,25,1)

solution1 = "10!/(2!8!) * 0.5^{10}"
solution2 = "10!/(1!9!) * 0.5^{10}"
solution3 = "1 * 0.5^{10}"
solution4 = "10!/(2!8!) * 0.5^{10}+ 10!/(1!9!) * 0.5^{10} + 1 * 0.5^{10}"
solution5 = "1- ( 10!/(2!8!) * 0.5^{10}+ 10!/(1!9!) * 0.5^{10} + 1 * 0.5^{10})"
solution6 = " (1- ( 10!/(2!8!) * 0.5^{{10}}+ 10!/(1!9!) * 0.5^{{10}} + 1 * 0.5^{{10}}) ) ^ {{{0}}}".format(a)
solution7 = " 1 -  ((1- ( 10!/(2!8!) * 0.5^{{10}}+ 10!/(1!9!) * 0.5^{{10}} + 1 * 0.5^{{10}}) ) ^ {{{0}}} )".format(a)

solutions = [solution1,solution2,solution3,solution4,solution5,solution6,solution7]
# hint text id
hint_text_id = ["cse103fall2016week1problem1part1","cse103fall2016week1problem1part2"]
    </script>

    <script type="text/javascript">
      //fill in number of hint
        var hint_number = 2;

      //fill in hint input box id
        var hintIds = ["input_57ccdc78397649aca59b419c95573254_2_1",
                "input_1ddc9e0c3eea4ba3b61b291a7cfb650d_2_1"];

      //fill in problem input box id
        var proIds = ["input_0463a3ce78dd46429e9e2fcf681616d9_2_1",
               "input_0463a3ce78dd46429e9e2fcf681616d9_3_1"];
               
      //fill in hint text id. Need to match the hint_text_id from above
        var hintTextIds = ["cse103fall2016week1problem1part1",
                  "cse103fall2016week1problem1part2"];
    </script>



<p>There is a class of $a students taking a T/F quiz of 10 questions. We would normally think it unlikely that a student would get more than 2 questions right out of 10 by flipping a coin, but with a class this large, it might be possible. Here we'll calculate the probability of at least one student in the class of size $a getting more than 2 questions right.</p><ul><li>What is the probability of a particular student getting exactly 2 questions right?</li></ul>

    <customresponse cfn="check" expect="\[$solution1\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<ul><li>What is the probability of a particular student getting exactly 1 question right?</li></ul>

    <customresponse cfn="check" expect="\[$solution2\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<ul><li>What is the probability of a particular student getting no questions right?</li></ul>

    <customresponse cfn="check" expect="\[$solution3\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<ul><li>What is the probability of a particular student getting at most 2 questions right?</li></ul>

    <customresponse cfn="check" expect="\[$solution4\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<ul><li>What is the probability of a particular student getting 3 or more questions right?</li></ul>

    <customresponse cfn="check" expect="\[$solution5\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<ul><li>What is the probability of every student in the class getting 3 or more questions right?</li></ul>

    <customresponse cfn="check" expect="\[$solution6\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<ul><li>What is the probability of some student in the class getting at most 2 questions right?</li></ul>

    <customresponse cfn="check" expect="\[$solution7\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>



    
    

    

    <script type="loncapa/python">
from hint import evaluate
from hint.get_hint import get_hint

# grader of the problem
def check(expect, ans):
  return evaluate.evaluate(expect, ans)

# hint function
def hint_fn(answer_ids, student_answers, new_cmap, old_cmap):
  aid = answer_ids[0]
  ans = str(student_answers[aid])
  i = int(aid[-3])-2

  if ans == "":
    return
  key = new_cmap.keys()[0]
  if new_cmap[key]['correctness'] == "correct":
    return

  hint = get_hint(solutions[i], ans, hint_text_id, i)
  new_cmap.set_hint_and_mode(aid,hint,'always')
    </script>

        
    <script type="text/javascript" src="/static/js/adaptive_hint.js"></script>
    <script type="text/javascript">                             
      for (var i=0; i != hint_number; i++) {
        hide_hint(i);
      }
      show_hint(hint_number);
    </script>

  </text>
</problem>