<problem>
  <text>
    <script type="loncapa/python">

# random variables (no need to import random library)

# Solutions with variables converted to string
# Make sure you name the solution with part id at the end. e.g. 'solution1' will be solution for part 1.
solution1 = "4*8"
solution2 = "(4 * 8) / C(52-5, 2)"


# Group all solutions into a list
solutions = [solution1, solution2]


# hint text id
hint_text_id = ["cse103fall2016week1problem1part1","cse103fall2016week1problem1part2"]
    </script>

    <script type="text/javascript">
      //fill in number of hint
        var hint_number = 2;

      //fill in hint input box id
        var hintIds = ["input_57ccdc78397649aca59b419c95573254_2_1",
                "input_1ddc9e0c3eea4ba3b61b291a7cfb650d_2_1"];

      //fill in problem input box id
        var proIds = ["input_0463a3ce78dd46429e9e2fcf681616d9_2_1",
               "input_0463a3ce78dd46429e9e2fcf681616d9_3_1"];
               
      //fill in hint text id. Need to match the hint_text_id from above
        var hintTextIds = ["cse103fall2016week1problem1part1",
                  "cse103fall2016week1problem1part2"];
    </script>



<p>Suppose you have been dealt 4\(\heartsuit\) and 5 \(\heartsuit\). What is the conditional probability that you will get a straight (five cards in sequence, but not all from the same suit) given that you have been dealt these two cards, and that the flop is "2\(\clubsuit\), Q\(\clubsuit\), K\(\diamondsuit\)"</p><p>Consider following events,</p><p>Event A: You will get a straight.</p><p>Event B:  You have been dealt with 5 cards i.e. 4\(\heartsuit\), 5 \(\heartsuit\), 2\(\clubsuit\), Q\(\clubsuit\), K\(\diamondsuit\).</p><p>In this case you need a 3 and either a 6 or A(ace) on the turn and river in order to get straight.</p><ul><li>The number of such card pairs, ignoring order, is the size of  \(|A \cap B| = \)</li></ul>

    <customresponse cfn="check" expect="\[$solution1\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<ul><li>The conditional probability \(P(A|B) =\)</li></ul>

    <customresponse cfn="check" expect="\[$solution2\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>



    
    

    

    <script type="loncapa/python">
from hint import evaluate
from hint.get_hint import get_hint

# grader of the problem
def check(expect, ans):
  return evaluate.evaluate(expect, ans)

# hint function
def hint_fn(answer_ids, student_answers, new_cmap, old_cmap):
  aid = answer_ids[0]
  ans = str(student_answers[aid])
  i = int(aid[-3])-2

  if ans == "":
    return
  key = new_cmap.keys()[0]
  if new_cmap[key]['correctness'] == "correct":
    return

  hint = get_hint(solutions[i], ans, hint_text_id, i)
  new_cmap.set_hint_and_mode(aid,hint,'always')
    </script>

        
    <script type="text/javascript" src="/static/js/adaptive_hint.js"></script>
    <script type="text/javascript">                             
      for (var i=0; i != hint_number; i++) {
        hide_hint(i);
      }
      show_hint(hint_number);
    </script>

  </text>
</problem>