<problem>
  <text>
    <script type="loncapa/python">


r = random.randrange(4,6,1)/10;#random(0.4,0.6,0.1);
k = random.randrange(100,200,1);

p = random.randrange(r*10+1,9,1)/10#random(r+0.1,0.9,0.1);
n2 = random.randrange(300,500,1);
k2 = random.randrange(251,299,1);
#ans1 = ceil(p*k/r);

solution1 = "{0}*{1}/{2}".format(p, k, r) #{Compute("$p*$k/$r")}
solution2 = "{0}*{1}/{2}".format(n2, r, k2) #{Compute("$n2*$r/$k2")->with(tolType=>'absolute', tolerance=>'0.01')}

solutions = [solution1, solution2]
# hint text id
hint_text_id = ["cse103fall2016week1problem1part1","cse103fall2016week1problem1part2"]
    </script>

    <script type="text/javascript">
      //fill in number of hint
        var hint_number = 2;

      //fill in hint input box id
        var hintIds = ["input_57ccdc78397649aca59b419c95573254_2_1",
                "input_1ddc9e0c3eea4ba3b61b291a7cfb650d_2_1"];

      //fill in problem input box id
        var proIds = ["input_0463a3ce78dd46429e9e2fcf681616d9_2_1",
               "input_0463a3ce78dd46429e9e2fcf681616d9_3_1"];
               
      //fill in hint text id. Need to match the hint_text_id from above
        var hintTextIds = ["cse103fall2016week1problem1part1",
                  "cse103fall2016week1problem1part2"];
    </script>



<h3>Markov's Inequality</h3><p>Markov's inequality relate probabilities to expectations, and provide bounds for the cumulative distribution function of a random variable. The Markov's inequality is stated as follow:</p><h2>If X is any nonnegative random variable and a &gt; 0, then  \[\mathbb{P}(X \geq a) \leq \frac{\mathbb{E}(X)}{a}\]</h2><p>John has a biased coin with \(P(\mathrm{heads}) = $r\). He tosses this coin \(N\) times and, out of the \(N\) times, the coin lands on heads $k times. Using Markov's inequality that he learned from CSE103, he says that the probability of seeing at least this many heads is at most $p.</p><p>(1) Give the best lower bound, using Markov inequality, on the number of times John tossed the coin?\(N\) =</p>

    <customresponse cfn="check" expect="\[$solution1\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<p>(Provide the exact answer, don't round it to the next larger integer number.)</p><p>Suppose John lends you this coin. If you flip the coin $n2 times, what is the upper bound of the probability of seeing at least $k2 heads using Markov's inequality?</p><p>(2) \(P(\mbox{Number of heads} \geq $k2)\) \(\le\)</p>

    <customresponse cfn="check" expect="\[$solution2\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>



    
    

    

    <script type="loncapa/python">
from hint import evaluate
from hint.get_hint import get_hint

# grader of the problem
def check(expect, ans):
  return evaluate.evaluate(expect, ans)

# hint function
def hint_fn(answer_ids, student_answers, new_cmap, old_cmap):
  aid = answer_ids[0]
  ans = str(student_answers[aid])
  i = int(aid[-3])-2

  if ans == "":
    return
  key = new_cmap.keys()[0]
  if new_cmap[key]['correctness'] == "correct":
    return

  hint = get_hint(solutions[i], ans, hint_text_id, i)
  new_cmap.set_hint_and_mode(aid,hint,'always')
    </script>

        
    <script type="text/javascript" src="/static/js/adaptive_hint.js"></script>
    <script type="text/javascript">                             
      for (var i=0; i != hint_number; i++) {
        hide_hint(i);
      }
      show_hint(hint_number);
    </script>

  </text>
</problem>