<problem>
  <text>
    <script type="loncapa/python">

# random variables (no need to import random library)

a = random.randint(1,5);
b = random.randint(6,10);

# Solutions with variables converted to string
# Make sure you name the solution with part id at the end. e.g. 'solution1' will be solution for part 1.
solution1 = "{0}*(1+1/2+1/3+1/4+1/5)".format(a)
solution2 = "{0}*(1+1/4+1/9+1/16+1/25)".format(b)
solution3 = "{0}*(1-1/2+1/3-1/4+1/5)".format(a)
solution4 = "{0}*(1+1/4+1/9+1/16+1/25)".format(b)
solution5 = "{0}*(1-2/2+1/3-2/4+1/5)".format(a)
solution6 = "{0}*(1+1/9+1/25) +{0}*4*(1/4+1/16)".format(b)


# Group all solutions into a list
solutions = [solution1,solution2,solution3,solution4,solution5,solution6]

# hint text id
hint_text_id = ["cse103fall2016week1problem1part1","cse103fall2016week1problem1part2"]
    </script>

    <script type="text/javascript">
      //fill in number of hint
        var hint_number = 2;

      //fill in hint input box id
        var hintIds = ["input_57ccdc78397649aca59b419c95573254_2_1",
                "input_1ddc9e0c3eea4ba3b61b291a7cfb650d_2_1"];

      //fill in problem input box id
        var proIds = ["input_0463a3ce78dd46429e9e2fcf681616d9_2_1",
               "input_0463a3ce78dd46429e9e2fcf681616d9_3_1"];
               
      //fill in hint text id. Need to match the hint_text_id from above
        var hintTextIds = ["cse103fall2016week1problem1part1",
                  "cse103fall2016week1problem1part2"];
    </script>



<p>Let \(X_1, X_2, \cdots X_{5}\) be independent random variables with:</p><p>\[\mathbb{E}(X_i) = \frac{$a}{i} \hspace{10pt} \mbox{VAR}(X_i) = \frac{$b}{i^2}\]</p><p>Compute:</p><ul><li>What is the mean of \(X_1 + X_2 + X_3 + X_4 + X_5\) ?</li></ul>

    <customresponse cfn="check" expect="\[$solution1\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<ul><li>How about the variance?</li></ul>

    <customresponse cfn="check" expect="\[$solution2\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<ul><li>What is the mean of \(X_1 - X_2 + X_3 - X_4 + X_5\) ?</li></ul>

    <customresponse cfn="check" expect="\[$solution3\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<ul><li>How about the variance?</li></ul>

    <customresponse cfn="check" expect="\[$solution4\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<ul><li>What is the mean of \(X_1 - 2X_2 + X_3 - 2X_4 + X_5\) ?</li></ul>

    <customresponse cfn="check" expect="\[$solution5\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<ul><li>How about the variance?</li></ul>

    <customresponse cfn="check" expect="\[$solution6\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>



    
    

    

    <script type="loncapa/python">
from hint import evaluate
from hint.get_hint import get_hint

# grader of the problem
def check(expect, ans):
  return evaluate.evaluate(expect, ans)

# hint function
def hint_fn(answer_ids, student_answers, new_cmap, old_cmap):
  aid = answer_ids[0]
  ans = str(student_answers[aid])
  i = int(aid[-3])-2

  if ans == "":
    return
  key = new_cmap.keys()[0]
  if new_cmap[key]['correctness'] == "correct":
    return

  hint = get_hint(solutions[i], ans, hint_text_id, i)
  new_cmap.set_hint_and_mode(aid,hint,'always')
    </script>

        
    <script type="text/javascript" src="/static/js/adaptive_hint.js"></script>
    <script type="text/javascript">                             
      for (var i=0; i != hint_number; i++) {
        hide_hint(i);
      }
      show_hint(hint_number);
    </script>

  </text>
</problem>