<problem>
  <text>
    <script type="loncapa/python">
showPartialCorrectAnswers = 1

nCards=random.randrange(4,9,1);
b=random.randrange(2,nCards-1,1);

solution1 = "C(52,{0})".format(nCards)
solution2 = "C(12,{0})".format(b) 
solution3 = "C(40,{0} - {1})".format(nCards,b)
solution4 = "C(12,{0}) * C(40,{1} - {0})".format(b,nCards) 
solution5 = "C(12,{0})*C(40,{1}-{0})/C(52,{1})".format(b,nCards)

solutions = [solution1,solution2,solution3,solution4,solution5]

# hint text id
hint_text_id = ["cse103fall2016week1problem1part1","cse103fall2016week1problem1part2"]
    </script>

    <script type="text/javascript">
      //fill in number of hint
        var hint_number = 2;

      //fill in hint input box id
        var hintIds = ["input_57ccdc78397649aca59b419c95573254_2_1",
                "input_1ddc9e0c3eea4ba3b61b291a7cfb650d_2_1"];

      //fill in problem input box id
        var proIds = ["input_0463a3ce78dd46429e9e2fcf681616d9_2_1",
               "input_0463a3ce78dd46429e9e2fcf681616d9_3_1"];
               
      //fill in hint text id. Need to match the hint_text_id from above
        var hintTextIds = ["cse103fall2016week1problem1part1",
                  "cse103fall2016week1problem1part2"];
    </script>



<h2>Probability of face cards</h2>
<p>A poker hand consisting of $nCards cards is dealt from a standard deck of 52 cards.
Find the probability that the hand contains exactly $b face cards (Jack, Queen or King).</p>
<p>First, we know the number of all possible hands of $nCards cards is</p>


    <customresponse cfn="check" expect="\[$solution1\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<p>Then, we calculate the number of hands that have exactly $b face cards.</p>
<p>The total number of face cards is 12 (Jack, Queen or King of all four suites). To choose $b from them we have</p>


    <customresponse cfn="check" expect="\[$solution2\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<p>possibilities.</p>
<p>The other $nCards-$b cards in the hand can be any of the non-face cards. There are a total of 52 - 12 = 40 non-face cards. To choose $nCards-$b form them, there are</p>


    <customresponse cfn="check" expect="\[$solution3\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<p>possibilities.</p>
<p>Thus we can compute the number of hands that have exactly $b face cards, which is</p>


    <customresponse cfn="check" expect="\[$solution4\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<p>Finally we can calculate the probability of such hands, by calculating the ratio of the number of such hands to the number of all hands. This is</p>


    <customresponse cfn="check" expect="\[$solution5\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


    

    

    <script type="loncapa/python">
from hint import evaluate
from hint.get_hint import get_hint

# grader of the problem
def check(expect, ans):
  return evaluate.evaluate(expect, ans)

# hint function
def hint_fn(answer_ids, student_answers, new_cmap, old_cmap):
  aid = answer_ids[0]
  ans = str(student_answers[aid])
  i = int(aid[-3])-2

  if ans == "":
    return
  key = new_cmap.keys()[0]
  if new_cmap[key]['correctness'] == "correct":
    return

  hint = get_hint(solutions[i], ans, hint_text_id, i)
  new_cmap.set_hint_and_mode(aid,hint,'always')
    </script>

        
    <script type="text/javascript" src="/static/js/adaptive_hint.js"></script>
    <script type="text/javascript">                             
      for (var i=0; i != hint_number; i++) {
        hide_hint(i);
      }
      show_hint(hint_number);
    </script>

  </text>
</problem>