# Stock category prediction using XGBoost

```python
import pickle
import pandas as pd
import numpy as np
import xgboost as xgb
from sklearn.model_selection import train_test_split
%pylab inline
```

## Read PCA parameters

```python
D=pickle.load(open('data/PCA_true.pickle','rb'))
col=D['columns']
eigvec=D['eigvec']
eigval=D['eigval']
```

### Sanity check
The following scatterplot should be very similar to the scatter plot you produced in notebook 3 for eigvec 1, eigvec 2 (indexing starts with 1)

```python
scatter(eigvec[:,0],eigvec[:,1])
```

## compute features

```python
#Taking the top 20 features(Eigen vectors)
features=eigvec[:,:20]
features.shape
```

##  Compute labels (sectors)

```python
TickerInfo=pd.read_csv('data/tickerInfo.tsv',sep='\t')
print(TickerInfo.shape)
TickerInfo.head()
```

## Creating necessary dictionaries

The Sectors dictionary below has the Sector name mapped to the sector ID. Using the Sectors dictionary below, create a dictionary mapping of the sector ID to indices incrementally as follows:

```
{
    'CD': 0,
    'CS': 1,
    'EN': 2,
    'FIN': 3,
    'HC': 4,
    'INDS': 5,
    'IT': 6,
    'MAT': 7,
    'RE': 8,
    'TS': 9,
    'UTIL': 10
}
```

In addition to this, you will need to create one more dictionary mapping index number to the sector name:
```
{
    0: 'Consumer Discretionary',
    1: 'Consumer Staples',
    2: 'Energy',
    3: 'Financials',
    4: 'Health Care',
    5: 'Industrials',
    6: 'Information Technology',
    7: 'Materials',
    8: 'Real Estate',
    9: 'Telecommunication Services',
    10: 'Utilities'
}
 ```
 
**<font color="magenta" size=2>Input</font>**
<br>
A dictionary **Sectors** as given in the following cell.

**<font color="blue" size=2>Output</font>**
<br>
Return two dictionaries **sector2number** and **number2sectorName** as mentioned in the description.

```python
Sectors={'Consumer Discretionary':'CD',
 'Consumer Staples':'CS',
 'Energy':'EN',
 'Financials':'FIN',
 'Health Care':'HC',
 'Industrials':'INDS',
 'Information Technology':'IT',
 'Materials':'MAT',
 'Real Estate':'RE',
 'Telecommunication Services':'TS',
 'Utilities':'UTIL'}

```

<!-- grade_id: cell-48e618d54073cb44, points: 0, grade: false, solution: true -->
```python
def get_sector_dicts(Sectors):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
sector2number, number2sectorName = get_sector_dicts(Sectors)
```

<!-- grade_id: cell-79897da0f4f488a4, points: 2, grade: true, solution: false -->
```python
# Hidden Tests here
```

```python
labels=[]
feature_vectors=[]
feature_vectors_test=[]
test_nos = []
for i in range(len(col)):
    c=col[i]
    if 'train' in c:
        ticker=c[6:-2]
        answer=list(TickerInfo[TickerInfo.Ticker==ticker]['SECTOR_ID'])
        if len(answer)==1:
            sector_no=sector2number[answer[0]]
            labels.append(sector_no)
            feature_vectors.append(features[i,:])
        else:
            print('error: could not find sector for ticker:',ticker)
    if 'test' in c:
        test_nos.append(c[5:-2])
        feature_vectors_test.append(features[i,:])
```

```python
len(labels),len(feature_vectors), len(test_nos),len(feature_vectors_test)
```

## XGBoost Setup

```python
X=np.array(feature_vectors)
y=np.array(labels)
X_test = np.array(feature_vectors_test)
y_test = np.array(test_nos)
X.shape, y.shape, X_test.shape, y_test.shape
```

```python
#Splitting between train and validation
X_train, X_valid, y_train, y_valid = train_test_split(X, y, test_size=0.30, random_state=6)
X_train.shape, X_valid.shape
```

```python
#Parameters
param = {}
param['max_depth']= 3   # depth of tree
param['eta'] = 0.3      # shrinkage parameter
param['silent'] = 1     # not silent
param['objective'] = 'multi:softmax'
param['nthread'] = 7 # Number of threads used
param['num_class']=11

num_round = 100
```

## Generating scores using XGBoost

The function <font color="blue">get_margin_scores</font> is used to predict the sector for each of the given samples.

**<font color="magenta" size=2>Input</font>**
<br>
1. **Training set** (X_train)
2. **Validation set** (X_valid)
3. **Training labels** (y_train)
4. **XGBoost Parameter List** (param)


**<font color="blue" size=2>Output</font>**
<br>
Return the following:
1. **y_pred_valid**: The raw output scores for the validation set


**Note**:
1. Round all raw scores to **three** decimal places
2. Remember to use **verbose_eval = False** while training and **ntree_limit=bst.best_ntree_limit** and **output_margin=True** while predicting
3. Remember to provide the **num_round** parameter while training and do not change it. We have currently set it to 100 (Refer previous cell). Not providing the parameter or changing it could produce different results.

<!-- grade_id: cell-0e8f383f79318be9, points: 0, grade: false, solution: true -->
```python
def get_margin_scores(X_train, X_valid, y_train, param):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
y_pred_valid = get_margin_scores(X_train, X_valid, y_train, param)
```

<!-- grade_id: cell-121c650f46b392f8, points: 0, grade: true, solution: false -->
```python
assert y_pred_valid.shape == (118, 11), "Incorrect shape"
assert type(y_pred_valid) == numpy.ndarray, "Incorrect type"
```

<!-- grade_id: cell-c9419db9f474f, points: 6, grade: true, solution: false -->
```python
# Hidden Tests Here
#If rounding off is the problem, give them 6 points out of 8
```

<!-- grade_id: cell-c9419ce7db9f474f, points: 2, grade: true, solution: false -->
```python
# Hidden Tests Here
```

## Computing Top1 and Top5 predictions

Using the margin scores generated, calculate top1 and top5 predictions for the given data:

**top1**: Find the most probable prediction for each example in the validation set

**top5**: Find the top 5 most probable predictions in descending order for each example in the validation set (most probable to fifth most probabale)

**<font color="magenta" size=2>Input</font>**
<br>
1. **Validation Output Scores** (y_pred_valid)

**<font color="blue" size=2>Output</font>**
<br>
1. **predictions_valid**: The most probable prediction for each example in the the validation set
2. **predictions_topn**: The top 5 predictions for each example in the validation set

**<font color="magenta" size=2>Sample input</font>**
1. **y_pred_valid**: [[-0.3, 1.2, 0.3, 0.5, -0.4, 0.0, 0.01, 1.0, -1.3, 0.2, -1.2], [0.4, -0.5, 1.3, -0.2, 0.6, -2.2, -0.8, 1.9, 0.9, -0.2, -1.7]]

**<font color="blue" size=2>Sample output</font>** 
1. **predictions_valid**: [1, 7]
2. **predictions_top5**: [[1,  7,  3,  2,  9], [ 7,  2,  8,  4,  0]]

<!-- grade_id: cell-12ee152fd08b1fdf, points: 0, grade: false, solution: true -->
```python
def get_predictions(y_pred_valid):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
predictions_valid, predictions_top5 = get_predictions(y_pred_valid)
```

<!-- grade_id: cell-02be6a89a5e44753, points: 0, grade: true, solution: false -->
```python
assert predictions_valid.shape == (118,), "Incorrect shape"
assert predictions_top5.shape == (118, 5), "Incorrect shape"
assert type(predictions_valid) == numpy.ndarray, "Incorrect type"
assert type(predictions_top5) == numpy.ndarray, "Incorrect type"
```

<!-- grade_id: cell-bd4d60a046981072, points: 3, grade: true, solution: false -->
```python
# Hidden Tests Here
```

<!-- grade_id: cell-96dc65236185c2de, points: 3, grade: true, solution: false -->
```python
# Hidden Tests Here
```

```python
acc = 0
for i in range(5):
    acc += sum(predictions_top5[:, i]==y_valid)
    print("Top ", i+1, ": \t", acc/len(y_valid))
```

## Generating the confusion matrix

What is a confusion matrix?
This is a useful link that explains this: https://en.wikipedia.org/wiki/Confusion_matrix#Example

We will now be using the top 2 values of the **predictions_top5** matrix generated, to produce the confusion matrix. We will be creating a confusion matrix by comparing the most probable prediction against the second most probable prediction. We are doing this to analyse the scenarios where the second most probable prediction is in fact the correct prediction. So remember to take **only the top2** from the **predictions_top5** matrix. For generating this matrix we do not need more than that.

An example with 4 classes (0, 1, 2, 3):

If 2 is getting confused with 1, i.e., if 1 is the top prediction and 2 is the second-top prediction, then your confusion matrix should add 1 to the cell (1, 2).

Different example scenarios:

**Scenario 1**:

**<font color="magenta" size=2>Sample Input</font>**<br>
<br>
y_label = 3<br>
top2_predictions = [2, 3]

**<font color="blue" size=2>Output</font>**<br>
confusion_matrix:<br>
[0 0 0 0]<br>
[0 0 0 0]<br>
[0 0 0 1]<br>
[0 0 0 0]<br>

**Scenario 2**:

**<font color="magenta" size=2>Sample Input</font>**<br>
<br>
Say we have 7 sample prediction values<br>
y_label = [3, 2, 3, 1, 0, 2, 2]<br>
top2_predictions = [[2, 3], [0, 3], [2, 3], [1, 2], [3, 0], [2, 0], [3, 2]]

**<font color="blue" size=2>Output</font>**<br><br>
confusion_matrix:<br>
[0 0 0 0]<br>
[0 0 0 0]<br>
[0 0 0 2]<br>
[1 0 1 0]<br>

**Explanation**:

The first example has top two predictions: (2, 3) and the label 3. So we add 1 to the position (2,3).

The second example has 7 sample predictions: 
1. In two scenarios 2 is predicted in place of 3 and the cell (2,3) is incremented twice.
2. In two other scenarios we have a case where the second prediction is right. The corresponding cells {(3,0) and (3,2)} are incremented once each.
3. In two scenarios, the first element is the correct prediction. No cell is incremented then.
4. In one scenario, neither of the top 2 predictions is correct. No cell is incremented then.

**Note**: y_label is the same as y_valid here.

<!-- grade_id: cell-8e31e89d9f64d4d6, points: 0, grade: false, solution: true -->
```python
def get_confusion_matrix(predictions_top5, y_valid):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
confusion_matrix = get_confusion_matrix(predictions_top5, y_valid)
```

<!-- grade_id: cell-157fe9078d64488b, points: 0, grade: true, solution: false -->
```python
assert confusion_matrix.shape == (11, 11), "Incorrect shape"
assert type(confusion_matrix) == numpy.ndarray, "Incorrect type"
assert type(confusion_matrix[0][0]) == numpy.int64, "Incorrect type"
```

<!-- grade_id: cell-b659fe7e4e082e91, points: 4, grade: true, solution: false -->
```python
# Hidden Tests Here
```

```python
for i in range(confusion_matrix.shape[0]):
    print("%30s" % number2sectorName[i], "\t", confusion_matrix[i, :])
```

## Interpretation of confusion matrix

Based on the confusion matrix generated, answer the following questions.

Categories:
```
    0) Consumer Discretionary
    1) Consumer Staples  
    2) Energy  
    3) Financials
    4) Health Care   
    5) Industrials   
    6) Information Technology   
    7) Materials   
    8) Real Estate   
    9) Telecommunication Services   
    10) Utilities  
```

Some standard instructions while answering the questions:
1. Each question has two parts: **ans** and **num_scen**
    1. **ans**: A list with numbers corresponding to the different categories. For example, if the answer is Consumer Discretionary and Consumer Staples, the **ans** should be a **list** containing **[0, 1]**. 
    2. **num_scen**: Number of scenarios where the condition in the given question is applicable. For example, if the condition given in the question happens 5 times, the **num_scen** should return **5** as an integer.
2. Type checks have been provided to validate the type of your answer
3. Remember that just the answer for the given matrix should be enough. You do not have to write a generic function that answers the question for all confusion matrices.

### Question 1

Which two sectors are most often confused with each other and how many times? 

```
    0) Consumer Discretionary  
    1) Consumer Staples  
    2) Energy  
    3) Financials
    4) Health Care   
    5) Industrials   
    6) Information Technology   
    7) Materials   
    8) Real Estate   
    9) Telecommunication Services   
    10) Utilities  
```

**Output format for this question**: The **ans1** must be list of tuples. If say sector 0 and sector 1 are confused with each other, then your answer must be [(0,1)]. If there are multiple pairs, then it must be a list of all the tuples. For example, **[(0,1),(2,3)]**.

The format for **num_scen1** is the same as mentioned in 1.9.

<!-- grade_id: cell-67fa65ee294bbf52, points: 0, grade: false, solution: true -->
```python
def question_1():
    # YOUR CODE HERE
    raise NotImplementedError()
    return ans1, num_scen1 
```

<!-- grade_id: cell-bf7603962387df88, points: 0, grade: true, solution: false -->
```python
ans1, num_scen1 = question_1()
assert type(ans1) == list, "Incorrect type"
assert type(ans1[0]) == tuple, "Incorrect type"
assert type(ans1[0][0]) == int, "Incorrect type"
assert type(num_scen1) == int, "Incorrect type"
```

<!-- grade_id: cell-d0a207dca2ce731e, points: 2, grade: true, solution: false -->
```python
# Hidden tests here
```

<!-- grade_id: cell-d0a20ca2ce731e, points: 2, grade: true, solution: false -->
```python
# Hidden tests here
```

### Question 2
Which sector(s) has/have the most number of scenarios where the prediction with the second highest score is actually correct? In other words, the sector receives the second highest score and but is correct. How many scenarios?

```
    0) Consumer Discretionary  
    1) Consumer Staples  
    2) Energy  
    3) Financials
    4) Health Care   
    5) Industrials   
    6) Information Technology   
    7) Materials   
    8) Real Estate   
    9) Telecommunication Services   
    10) Utilities  
```

**Note**: Output format for this question is the same as the output format mentioned in Section 1.9.

<!-- grade_id: cell-34137c196c56dd11, points: 0, grade: false, solution: true -->
```python
def question_2():
    # YOUR CODE HERE
    raise NotImplementedError()
    return ans2, num_scen2 
```

<!-- grade_id: cell-3a56a331a146b70f, points: 0, grade: true, solution: false -->
```python
ans2, num_scen2 = question_2()
assert type(ans2) == list, "Incorrect type"
assert type(ans2[0]) == int, "Incorrect type"
assert type(num_scen2) == int, "Incorrect type"
```

<!-- grade_id: cell-d0a7dca2ce731e, points: 1, grade: true, solution: false -->
```python
# Hidden tests here
```

<!-- grade_id: cell-d0a20ca2ce1e, points: 1, grade: true, solution: false -->
```python
# Hidden tests here
```

### Question 3

Which sector(s) most often identified incorrectly? In other words, the sector recieves the highest score even though the sector with the second highest score is the correct sector. How many times does this happen for each of these sectors?
```
    0) Consumer Discretionary  
    1) Consumer Staples  
    2) Energy  
    3) Financials
    4) Health Care   
    5) Industrials   
    6) Information Technology   
    7) Materials   
    8) Real Estate   
    9) Telecommunication Services   
    10) Utilities  
```

**Note**: Output format for this question is the same as the output format mentioned in Section 1.9.

<!-- grade_id: cell-34137c196c5, points: 0, grade: false, solution: true -->
```python
def question_3():
    # YOUR CODE HERE
    raise NotImplementedError()
    return ans3, num_scen3 
```

<!-- grade_id: cell-3aa146b70f, points: 0, grade: true, solution: false -->
```python
ans3, num_scen3 = question_3()
assert type(ans3) == list, "Incorrect type"
assert type(ans3[0]) == int, "Incorrect type"
assert type(num_scen3) == int, "Incorrect type"
```

<!-- grade_id: cell-d0a2ce731e, points: 1, grade: true, solution: false -->
```python
# Hidden tests here
```

<!-- grade_id: cell-d0a20c1e, points: 1, grade: true, solution: false -->
```python
# Hidden tests here
```

### Note for question 4 and 5
**Note**: The next set of questions might require you to generate a different confusion matrix. Feel free to change the original function or use the box below to write a new function for the same.

<!-- grade_id: cell-81118309a7120376, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

### Question 4

Which category/categories is/are never identified? In other words, which category/categories are never predicted as the top prediction?
```
    0) Consumer Discretionary  
    1) Consumer Staples  
    2) Energy  
    3) Financials
    4) Health Care   
    5) Industrials   
    6) Information Technology   
    7) Materials   
    8) Real Estate   
    9) Telecommunication Services   
    10) Utilities  
```

**Note**: 
1. This question does not have a second part, so your function has to return just one answer.
2. Output format for this question is the same as the output format mentioned in Section 1.9.

<!-- grade_id: cell-34137c1c5, points: 0, grade: false, solution: true -->
```python
def question_4():
    # YOUR CODE HERE
    raise NotImplementedError()
    return ans4
```

<!-- grade_id: cell-3aa146b7lkbn0f, points: 0, grade: true, solution: false -->
```python
ans4 = question_4()
assert type(ans4) == list, "Incorrect type"
assert type(ans4[0]) == int, "Incorrect type"
```

<!-- grade_id: cell-d0a2ce7jn31e, points: 2, grade: true, solution: false -->
```python
# Hidden tests here
```

### Question 5

Which sector(s) is/are most often missed while classifying? In other words, find the sector for which there is the largest number of examples such that the correct label does not appear as the top prediction.


```
    0) Consumer Discretionary  
    1) Consumer Staples  
    2) Energy  
    3) Financials
    4) Health Care   
    5) Industrials   
    6) Information Technology   
    7) Materials   
    8) Real Estate   
    9) Telecommunication Services   
    10) Utilities  
```

**Note**: Output format for this question is the same as the output format mentioned in Section 1.9.

<!-- grade_id: cell-34vd137c1c5, points: 0, grade: false, solution: true -->
```python
def question_5():
    # YOUR CODE HERE
    raise NotImplementedError()
    return ans5, num_scen5 
```

<!-- grade_id: cell-3aa146bsdv7lkbn0f, points: 0, grade: true, solution: false -->
```python
ans5, num_scen5 = question_5()
assert type(ans5) == list, "Incorrect type"
assert type(num_scen5) == int, "Incorrect type"
```

<!-- grade_id: cell-d0sda2ce7jn31e, points: 2, grade: true, solution: false -->
```python
# Hidden tests here
```

<!-- grade_id: cell-d0vdahb20c1e, points: 2, grade: true, solution: false -->
```python
# Hidden tests here
```

## Test set

List the pair of top two sectors for each test ticker.
Based on your validation results, estimate:
1. ** Accuracy1:** What is the frequency in which the correct sector is the first element in the pair.
2. ** Accuracy2:** What is the frequency in which the correct sector is in the pair.

### Generating test scores

The function <font color="blue">get_margin_scores_test</font> is used to predict the sector for each of the given test samples. Split the input data into train and validation in different ways and average the prediction scores over a number of iterations. You can experiment with the number, but you will need to ensure you keep to the time limit. (This should not take you more than a minute)

We estimate the predictions on the validation and the test set, compute the accuracy on the validation set and report the test predictions with the validation accuracy. In this scenario, we report the average top1 and top2 accuracy.

**<font color="magenta" size=2>Input</font>**
<br>
1. **Input data** (X)
2. **Test data** (X_test)
3. **Input labels** (y)
4. **XGBoost Parameter List** (param)


**<font color="blue" size=2>Output</font>**
<br>
Return the following:
1. **y_pred_test**: The raw output scores for the test set
2. **top1_acc**: Top 1 accuracy on the validation set
3. **top2_acc**: Top 2 accuracy on the validation set


**Note**:
1. You can reuse/call the <font color="blue">get_margin_scores</font> function or rewrite it in this function.
2. Instructions for the <font color="blue">get_margin_scores</font> apply for this too.

<!-- grade_id: cell-396107076b2ccef8, points: 0, grade: false, solution: true -->
```python
def get_margin_scores_test(X, X_test, y, param):
    # YOUR CODE HERE
    raise NotImplementedError()
```

```python
y_pred_test, top1_acc, top2_acc = get_margin_scores_test(X, X_test, y, param)
```

<!-- grade_id: cell-4c7bb5430759629b, points: 0, grade: true, solution: false -->
```python
assert type(y_pred_test) == np.ndarray, ""
assert y_pred_test.shape == (89, 11)
```

<!-- grade_id: cell-b877235a2a80c672, points: 2, grade: true, solution: false -->
```python
# Hidden tests here
```

<!-- grade_id: cell-0329170a30973225, points: 2, grade: true, solution: false -->
```python
# Hidden tests here
```

<!-- grade_id: cell-38b270851d859ac7, points: 2, grade: true, solution: false -->
```python
# Hidden tests here
```

### Computing the sector predictions

**Note**: This section is not evaluated. This is merely to see your predictions on the test set and report the top1 and top2 accuracy

```python
predictions1, predictions5 = get_predictions(y_pred_test) #This line will fail if get_predictions is not defined
i=0
print("%10s" % "Test ID", "%30s" % "Top prediction", "%30s" % "Second Top prediction")
for entry in predictions5[:, :2]:
    print("%10s" % test_nos[i], "%30s" % number2sectorName[entry[0]], "%30s" % number2sectorName[entry[1]])
    i += 1
top1_acc_perc = top1_acc*100
top2_acc_perc = top2_acc*100
print("\n\n\nTop 1 accuracy: %s" % top1_acc_perc, "%")
print("\nTop 2 accuracy: %s" % top2_acc_perc, "%")
```