<font size="4" style="color:red;"> **IMPORTANT: ** Only modify cells which have the following comment</font>

```python
# modify this cell
```

<font style="color:red;"> Do not add any new cells when submitting homework. To test out new code, use the coding **scratchpad** by clicking the triangular icon in the bottom right corner of the screen. (**hotkey:** control-B)  </font>



# Regression I

This notebook is a an intuitive introduction to the concept of linear regression. Partially based on Chapter 10 in the book [**Statistics** by Freedman, Pisani and Purves](https://www.amazon.com/Statistics-4th-David-Freedman/dp/0393929728) 

```python
%pylab inline
import pandas as pd
import numpy as np
from numpy import arange,array,ones,linalg
```

## Small example
Suppose we have 9 points on the plane, defined by their $(x,y)$ coordinates

```python
x = arange(0,9)
y = [19, 20, 20.5, 21.5, 22, 23, 23, 25.5, 24]
plot(x,y,'bo');
```

* These points are close to a straight line trending upwoards. We would like to find that line. 

* The line is described by a function of the form: $f_{w_0,w_1}(x) \doteq w_0+w_1 x$

* our goal is to find $w_0$ and $w_1$.

* In a previous notebook we saw how to find a line that passes through 2 points.

* As here we have more than two points the system is **overconstrained** - there is no straight line that passes through all of the points.

* The points do not fall **exactly** on a line. So we are looking for $w_0,w_1$ 
such that the line is **closest** to the points.

We define the **Square difference** between the line $(w_0,w_1)$ and the points
$\langle (x_1,y_1),(x_2,y_2),\ldots,(x_n,y_n) \rangle$ to be 
$$\sum_{i=1}^n [f_{w_0,w_1}(x_i)-y_i]^2$$
(In our example $n=9$)

The values of $w_0,w_1$ which minimize the square difference, are called the **least squares** solution.

We use `numpy.linalg` to find the minimum. 

To define the problem for `numpy.linalg` we use matrix notation. We define ${\bf A}$ to be the $n$ by $2$ matrix:
$${\bf A} =  \begin{pmatrix}
 1, x_1\\
 \vdots\\
 1, x_n
 \end{pmatrix} 
$$

 ${\bf y}$ and ${\bf w}$ to be column vectors:
 $${\bf y} =  \begin{pmatrix}
 y_1\\
 \vdots\\
 y_n
 \end{pmatrix} 
 ,\;\;\;\;
 {\bf w} = \begin{pmatrix}
 w_0\\
 w_1
 \end{pmatrix}
$$

We can then express the differences as a vector ${\bf d}$:
$$ {\bf d} = {\bf A} {\bf w} - {\bf y} $$

The square difference, that we wish to minimize, is $\|{\bf d}\|_2^2 = \sum_{i=1}^n d_i^2$

```python
A = array([ ones(9),x ])
# linearly generated sequence
w = linalg.lstsq(A.T,y)[0] # obtaining the parameters
w
```

```python
# plotting the line and the square differences
line = w[0]+w[1]*x # regression line
plot(x,line,'r-',x,y,'o')

for i in range(len(x)):
    plot([x[i],x[i]],[y[i],w[1]*x[i]+w[0]],'g')
```

## A real-life example
We now turn our attention to the analysis of some real-word data. This data describes  the weight and height of some 25,000 people.

Weight/Weight data downloaded from [here](https://people.sc.fsu.edu/~jburkardt/data/csv/csv.html)

```python
HW=pd.read_csv('../data/HW25000.csv')
HW=HW.iloc[:,1:]
HW.columns=['Height','Weight']
HW.head()
```

```python
# Copute some basic descritive statistics.
HW.describe()
```

```python
A=np.array(HW['Height'])
A=np.array([ones(len(A)),A])
y=np.array(HW['Weight'])
A.shape,y.shape

w1 = linalg.lstsq(A.T,y)[0] # finding the optimal parameters
w1
```

```python
def f(x,w):
    return w[0]+w[1]*x
```

```python
ax=HW.plot(kind='scatter',s=1,x='Height',y='Weight',figsize=[10,8]);
x0,x1=xlim()
ax.plot([x0,x1],[f(x0,w1),f(x1,w1)],'r');
```

## The graph of averages 

```python
# calculate the mean weight for each 1-inch interval of height.
HW['round_height']=HW['Height'].round()
HW['round_weight']=HW['Weight'].round()

per_height_means=HW.groupby('round_height').mean()[['Weight']]
```

```python
ax=HW.plot(kind='scatter',s=1,x='Height',y='Weight',figsize=[10,8]);
per_height_means.plot(y='Weight',style='ro',ax=ax,legend=False)

_xlim=xlim()
_ylim=ylim()
for _x in arange(_xlim[0]+0.5,_xlim[1],1):
    plot([_x,_x],[_ylim[0],_ylim[1]],'g')

x0,x1=xlim()
ax.plot([x0,x1],[f(x0,w1),f(x1,w1)],'k');
```

## The two regression lines

```python
A=np.array(HW['Weight'])
A=np.array([ones(len(A)),A])
y=np.array(HW['Height'])
A.shape,y.shape

w2 = linalg.lstsq(A.T,y)[0] # finding the optimal parameters
w2
```

```python
ax=HW.plot(kind='scatter',s=1,x='Height',y='Weight',figsize=[10,8]);
x0,x1=xlim()
ax.plot([x0,x1],[f(x0,w1),f(x1,w1)],'r');

y0,y1=ylim()
ax.plot([f(y0,w2),f(y1,w2)],[y0,y1],'k');

```

# Excercises

It is now your turn to try out all of the above code. For this exercise you will measure the relationship between hours studyig and student grades for a simulated (*fake*) dataset. Below is a scatter plot of the data.

```python
data = pd.read_csv('../data/hw/hw_regression_data.csv')
print data.shape
data.head()
```

```python
ax= data.plot(kind='scatter',s=1.4,x="study_hours",y="grades",figsize=[10,8]);
```

## Exercise 1: Get Averages

Write a function, **get_averages**, that:
1. For all data points round the student's study hours to the nearest whole number.
2. Compute the mean grade for each rounded whole number of study hours.

The function should return a pandas dataframe of the mean grades for each existing rounded whole number. The beginning of the panda dataframe is provided for you below:

<font  style="color:blue"> * **Code:** *</font>
```python
grade_means = get_averages(data)

print type( grade_means )
print grade_means.shape
print grade_means.columns 
grade_means.head()
```


<font  style="color:magenta"> * **Output:** *</font>
```
<class 'pandas.core.frame.DataFrame'>
(38, 1)
Index([u'grades'], dtype='object')
```
<img src="../data/hw/reg_ex_table.png"  style="width: 190px;"/>

<!-- grade_id: cell-3db08936b836b64c, points: 0, grade: false, solution: true -->
```python
# modify this cell

def get_averages(data):
    # input: the HW's dataset
    # output: a pandas dataframe yielding the mean grade for each rounded number of study hours
    
    ### BEGIN SOLUTION
    data['round_study_hours']=data['study_hours'].round()
    return data.groupby('round_study_hours').mean()[['grades']]
    ### END SOLUTION
```

<!-- grade_id: cell-bd0897e089d912f5, points: 1, grade: true, solution: false -->
```python
# check student function
grade_means = get_averages(data)

assert grade_means.shape == (38, 1)
assert str(type(grade_means)) == "<class 'pandas.core.frame.DataFrame'>"

try: assert grade_means.columns[0] == "grades"
except: raise ValueError('the column in "grade_means" must be named grades')

try: assert sum(abs( grade_means.tail()["grades"].values - \
               np.array([ 90.77739,90.02272,92.97258,91.87963,93.48365])    )) < 10**-4
except: raise ValueError('some values are incorrect in grade_means')
 

# plot Graph
ax= data.plot(kind='scatter',s=1.4,x="study_hours",y="grades",figsize=[10,8]);
grade_means.plot(y='grades',style='ro', ax=ax,legend=False)
_xlim = xlim()
_ylim =ylim()
for _x in arange(_xlim[0]+0.5,_xlim[1],1):
    plot([_x,_x],[_ylim[0],_ylim[1]],'g')

### BEGIN HIDDEN TESTS
assert sum(abs( grade_means["grades"][20:30].values - \
        array([64.74628137,65.6715546,65.25890137,68.44169892,67.90134068,70.61057739,69.79446048,  \
               71.09018443,71.72610072,72.67667563,72.63601539])   )) < 5*10**-4
### BEGIN HIDDEN TESTS

```

## Exercise 2: Simple Linear Regression

Write a function, **do_regression**, that performs linear regression to return 
$ \;
 {\bf w} = \begin{pmatrix}
 w_0\\
 w_1
 \end{pmatrix}
$
for the equation:

$$
\text{grades}= w_0 + w_1 \cdot \text{study_hours}
$$


<font  style="color:blue"> * **Code:** *</font>
```python
w = do_regression(data)
print type(w)
print w.shape
```

<font  style="color:magenta"> * **Output:** *</font>
```
<type 'numpy.ndarray'>
(2,)
```


<!-- grade_id: cell-984f5b96264b6d11, points: 0, grade: false, solution: true -->
```python
# modify this cell

def do_regression(data):
    # input: the HW's dataset
    # output: a numpy array yielding w=(w0,w1) from linear regression
    
    ### BEGIN SOLUTION
    A=np.array(data['study_hours'])
    A=np.array([ones(len(A)),A])
    y=np.array(data['grades'])
    return linalg.lstsq(A.T,y)[0]
    ### END SOLUTION
```

<!-- grade_id: cell-db30bb6b08376d64, points: 1, grade: true, solution: false -->
```python
# check student function
w1_hw = do_regression(data)
assert str(type(w1_hw)) ==  "<type 'numpy.ndarray'>"
assert w1_hw.shape == (2,)
assert abs(w1_hw[0] - 48.0694704904) < 10**-5

# plot Means
ax= data.plot(kind='scatter',s=1.4,x="study_hours",y="grades",figsize=[10,8]);
x0,x1=xlim()
ax.plot([x0,x1],[f(x0,w1_hw),f(x1,w1_hw)],'k');

### BEGIN HIDDEN TESTS
abs(w1_hw[1] - 0.840093766443) < 10**-5
### BEGIN HIDDEN TESTS

```

## Exercise 3: Reversed Regression

Write a function, **reverse_regression**, that again performs linear regression to return 
$ \;
 {\bf w} = \begin{pmatrix}
 w_0\\
 w_1
 \end{pmatrix}
$
but this time for:

$$
\text{study_hours}= w_0 + w_1 \cdot \text{grades}
$$


<font  style="color:blue"> * **Code:** *</font>
```python
w = reverse_regression(data)
print type(w)
print w.shape
```

<font  style="color:magenta"> * **Output:** *</font>
```
<type 'numpy.ndarray'>
(2,)
```


<!-- grade_id: cell-3487ea247243635d, points: 0, grade: false, solution: true -->
```python
# modify this cell

def reverse_regression(data):
    # input: the HW's dataset
    # output: a numpy array yielding w=(w0,w1) for the reversed linear regression
    
    ### BEGIN SOLUTION
    A=np.array(data['grades'])
    A=np.array([ones(len(A)),A])
    y=np.array(data['study_hours'])
    return linalg.lstsq(A.T,y)[0] 
    ### END SOLUTION

```

<!-- grade_id: cell-2dfb7052d60f8919, points: 1, grade: true, solution: false -->
```python
# check answer
w2_hw = reverse_regression(data)
assert str(type(w2_hw)) ==  "<type 'numpy.ndarray'>"
assert w2_hw.shape == (2,)
assert abs(w2_hw[0] - -31.85141332) < 10**-5

# plot graph
ax= data.plot(kind='scatter',s=1.4,x="study_hours",y="grades",figsize=[10,8]);
x0,x1=xlim()
ax.plot([x0,x1],[f(x0,w1_hw),f(x1,w1_hw)],'k');
y0,y1=ylim()
ax.plot([f(y0,w2_hw),f(y1,w2_hw)],[y0,y1],'r');

### BEGIN HIDDEN TESTS
abs(w2_hw[1] - 0.86337634) < 10**-5
### END HIDDEN TESTS
```

```python





```