<problem>
  <text>
    <script type="loncapa/python">

x=[0]*4
m=[0]*4

for i in xrange(4):
 	x[i] = random.randrange(1,9,1)
	m[i] = random.randrange(1,12,1)/2.0
prob = [[0 for i in xrange(x[j]+1)] for j in xrange(4)]

ans0 = "{0}^{{{1}}}*e^{{-{0}}}/{1}!".format(m[0],x[1])

for i in xrange(4):
	for j in xrange(x[i]+1):
		prob[i][j] = "{0}^{{{1}}}/e^{{{0}}}/{1}!".format(m[i],j)

sum1 = '+'.join(prob[1])
# sum1 = ''
# for j in xrange(x[1]+1):
# 	sum1 = '{}+{}'.format(sum1, prob[1][j])

sum2 = '+'.join(prob[2])
ans2 = '1-(' + sum2 + ')'
# sum2 = ''
# for j in xrange(x[2]+1):
# 	sum2 = "({}+{})".format(sum2,prob[2][j])
# 	ans2 = "(1-{})".format(sum2)

sum3 = '+'.join(prob[3])
# sum3 = ''
# for j in xrange(x[3]):
# 	sum3 = sum3+prob[3][j]


x0 = x[0]
x1 = x[1]
x2 = x[2]
x3 = x[3]
m0 = m[0]
m1 = m[1]
m2 = m[2]
m3 = m[3]

solution1 = ans0
solution2 = sum1
solution3 = ans2
solution4 = sum3

solutions=[solution1, solution2, solution3, solution4]

# hint text id
hint_text_id = ["cse103fall2016week1problem1part1","cse103fall2016week1problem1part2"]
    </script>

    <script type="text/javascript">
      //fill in number of hint
        var hint_number = 2;

      //fill in hint input box id
        var hintIds = ["input_57ccdc78397649aca59b419c95573254_2_1",
                "input_1ddc9e0c3eea4ba3b61b291a7cfb650d_2_1"];

      //fill in problem input box id
        var proIds = ["input_0463a3ce78dd46429e9e2fcf681616d9_2_1",
               "input_0463a3ce78dd46429e9e2fcf681616d9_3_1"];
               
      //fill in hint text id. Need to match the hint_text_id from above
        var hintTextIds = ["cse103fall2016week1problem1part1",
                  "cse103fall2016week1problem1part2"];
    </script>



<p>Given that \(X\) is a random variable having a Poisson distribution, compute the following:</p>
<p>(a) When \(\mu = $m0 \), \( P(X = $x0) = \)</p>


    <customresponse cfn="check" expect="\[$solution1\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<p>(b) When \(\mu = $m1 \), \( P(X \leq $x1) = \)</p>


    <customresponse cfn="check" expect="\[$solution2\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<p>(c) When \(\mu = $m2 \), \( P(X &gt; $x2) = \)</p>


    <customresponse cfn="check" expect="\[$solution3\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


<p>(d) When \(\mu = $m3 \), \( P(X &lt; $x3) = \)</p>


    <customresponse cfn="check" expect="\[$solution4\]">
      <textline/>
      <hintgroup hintfn="hint_fn"/>
    </customresponse>


    

    

    <script type="loncapa/python">
from hint import evaluate
from hint.get_hint import get_hint

# grader of the problem
def check(expect, ans):
  return evaluate.evaluate(expect, ans)

# hint function
def hint_fn(answer_ids, student_answers, new_cmap, old_cmap):
  aid = answer_ids[0]
  ans = str(student_answers[aid])
  i = int(aid[-3])-2

  if ans == "":
    return
  key = new_cmap.keys()[0]
  if new_cmap[key]['correctness'] == "correct":
    return

  hint = get_hint(solutions[i], ans, hint_text_id, i)
  new_cmap.set_hint_and_mode(aid,hint,'always')
    </script>

        
    <script type="text/javascript" src="/static/js/adaptive_hint.js"></script>
    <script type="text/javascript">                             
      for (var i=0; i != hint_number; i++) {
        hide_hint(i);
      }
      show_hint(hint_number);
    </script>

  </text>
</problem>